Positive selection and relaxed purifying selection contribute to rapid evolution of male-biased genes in a dioecious flowering plant

Author Response

The following is the authors’ response to the original reviews.

Thank you for reviewing our manuscript. We do find that the reviews are constructive and meaningful. Accordingly, we incorporated most suggestions into our revision. We provided a point-by-point responses to the reviews below.

Reviewer #1 (Public Review):

The evolution of dioecy in angiosperms has significant implications for plant reproductive efficiency, adaptation, evolutionary potential, and resilience to environmental changes. Dioecy allows for the specialization and division of labor between male and female plants, where each sex can focus on specific aspects of reproduction and allocate resources accordingly. This division of labor creates an opportunity for sexual selection to act and can drive the evolution of sexual dimorphism.

In the present study, the authors investigate sex-biased gene expression patterns in juvenile and mature dioecious flowers to gain insights into the molecular basis of sexual dimorphism. They find that a large proportion of the plant transcriptome is differentially regulated between males and females with the number of sex-biased genes in floral buds being approximately 15 times higher than in mature flowers. The functional analysis of sex-biased genes reveals that chemical defense pathways against herbivores are up-regulated in the female buds along with genes involved in the acquisition of resources such as carbon for fruit and seed production, whereas male buds are enriched in genes related to signaling, inflorescence development and senescence of male flowers. Furthermore, the authors implement sophisticated maximum likelihood methods to understand the forces driving the evolution of sexbiased genes. They highlight the influence of positive and relaxed purifying selection on the evolution of male-biased genes, which show significantly higher rates of nonsynonymous to synonymous substitutions than female or unbiased genes. This is the first report (to my knowledge) highlighting the occurrence of this pattern in plants. Overall, this study provides important insights into the genetic basis of sexual dimorphism and the evolution of reproductive genes in Cucurbitaceae.

Thank you for your positive comments. Greatly appreciated.

There are, however, parts of the manuscript that are not clearly described or could be otherwise improved.

The number of denovo-assembled unigenes seems large and I would like to know how it compares to the number of genes in other Cucurbitaceae species. The presence of alternatively assembled isoforms or assembly artifacts may be still high in the final assembly and inflate the numbers of identified sex-biased genes.

The majority of unigenes were annotated by homologs in species of Cucurbitaceae (63%), including Momordica charantia (16.3%), Cucumis melo (11.9%), Cucurbita pepo (11.9%), Cucurbita moschata (11.5%), Cucurbita maxima (10.1%) and other species of Cucurbitaceae (Fig. S1C). We admit that in the final assembly, transcripts may be still overestimated due to the unavoidable presence of isoforms, although we have tried our best to filter it by several strategies of clustering methods. Additionally, we assessed the transcripts using BUSCOv5.4.5 and embryophyta_odb10 database with 1,614 plant orthologs assessment. Some 95.0% of these orthologs were covered by the unigenes, in which 1447 (89.7%) BUSCO genes were “Complete BUSCOs”, 85 (5.3%) were “Fragmented BUSCOs”, and only 82 (5.0%) were “Missing BUSCOs” (Table S2). Overall, our assessment suggested that we have generated high-quality reference transcriptomes in the absence of a reference genome. Subsequently, we revised the manuscript (lines 175-181).

It is interesting that the majority of sex-biased genes are present in the floral buds but not in the mature flowers. I think this pattern could be explored in more detail, by investigating the expression of male and female sex-biased genes throughout the flower development in the opposite sex. It is also not clear how the expression of the sex-biased genes found in the buds changes when buds and mature flowers are compared within each sex.

Thank you for your advice for further understanding of this interesting pattern. In the near future, we would like to study these issues through more development stages of flowers in each sex, probably with the aid of single-cell techniques and a reference genome. We have revised the manuscript to reflect these in Results, in the section "Tissue-biased/stage-biased gene expression" (lines 202216).

The statistical analysis of evolutionary rates between male-biased, female-biased, and unbiased genes is performed on samples with very different numbers of observations, therefore, a permutation test seems more appropriate here.

Thank you for your suggestion. However, all comparisons between sex-biased and unbiased genes were tested using Wilcoxon rank sum test in R software, which is more commonly used. Additionally, we tested some datasets, which were consistent with Wilcoxon rank sum test.

The impact of pleiotropy on the evolutionary rates of male-biased genes is speculative since only two tissue samples (buds and mature flowers) are used. More tissue types need to be included to draw any meaningful conclusions here.

Thank you for your advice for further understanding of the impact of pleitropy. In the near future, we would like make further investigations through more development stages of flowers and new technologies in each sex to consolidate the conclusion.

Reviewer #2 (Public Review):

Summary:

This study uses transcriptome sequence from a dioecious plant to compare evolutionary rates between genes with male- and female-biased expression and distinguish between relaxed selection and positive selection as causes for more rapid evolution. These questions have been explored in animals and algae, but few studies have investigated this in dioecious angiosperms, and none have so far identified faster rates of evolution in male-biased genes (though see Hough et al. 2014 https://doi.org/10.1073/pnas.1319227111).

Strengths:

The methods are appropriate to the questions asked. Both the sample size and the depth of sequencing are sufficient, and the methods used to estimate evolutionary rates and the strength of selection are appropriate. The data presented are consistent with faster evolution of genes with male-biased expression, due to both positive and relaxed selection.

This is a useful contribution to understanding the effect of sex-biased expression in genetic evolution in plants. It demonstrates the range of variation in evolutionary rates and selective mechanisms, and provides further context to connect these patterns to potential explanatory factors in plant diversity such as the age of sex chromosomes and the developmental trajectories of male and female flowers.

Weaknesses:

The presence of sex chromosomes is a potential confounding factor, since there are different evolutionary expectations for X-linked, Y-linked, and autosomal genes. Attempting to distinguish transcripts on the sex chromosomes from autosomal transcripts could provide additional insight into the relative contributions of positive and relaxed selection.

Thank you for your meanful suggestions. We agree that the identification of chromosome origins for transcripts would greatly improve the insights of selection, and we will investigate these issues, probably with a reference genome in the near future.

Reviewer #3 (Public Review):

The potential for sexual selection and the extent of sexual dimorphism in gene expression have been studied in great detail in animals, but hardly examined in plants so far. In this context, the study by Zhao, Zhou et al. al represents a welcome addition to the literature.

Relative to the previous studies in Angiosperms, the dataset is interesting in that it focuses on reproductive rather than somatic tissues (which makes sense to investigate sexual selection), and includes more than a single developmental stage (buds + mature flowers).

The main limitation of the study is the very low number of samples analyzed, with only three replicate individuals per sex (i.e. the whole study is built on six individuals only). This provides low power to detect differential expression. Along the same line, only three species were used to evaluate the rates of non-synonymous to synonymous substitutions, which also represents a very limited dataset, in particular when trying to fit parameter-rich models such as those implemented here.

A third limitation relates to the absence of a reference genome for the species, making the use of a de novo transcriptome assembly necessary, which is likely to lead to a large number of incorrectly assembled transcripts. Of course, the production of a reference transcriptome in this non-model species is already a useful resource, but this point should at least be acknowledged somewhere in the manuscript.

Each of these shortcomings is relatively important, and together they strongly limit the scope of the conclusions that can be made, and they should at least be acknowledged more prominently. The study is valuable in spite of these limitations and the topic remains grossly understudied, so I think the study will be of interest to researchers in the field, and hopefully inspire further, more comprehensive analyses.

We acknowledged that our sample size was relatively small. We will investigate these issues at the population level, probably with a reference genome in the near future. We acknowledged in the revised manuscript that there may be some incorrectly assembled transcripts. We assessed the transcripts using BUSCOv5.4.5 and the latest embryophyta_odb10 database with 1,614 plant orthologs assessment. As mentioned, 95.0% of these orthologs were covered by the unigenes, which of 1447 (89.7%) BUSCO genes were “Complete BUSCOs”, 85 (5.3%) were “Fragmented BUSCOs”, and only 82 (5.0%) were “Missing BUSCOs” (Table S2). In short, the quality of transcriptome was high in the absence of a reference genome.

Reviewer #1 (Recommendations For The Authors):

My main criticism of this manuscript is that it refers to gene names and orthogroups throughout the text, however, the assembled transcripts are not accessible. The reference trascriptome, orthology data, and alignments used for evolutionary analysis should be made available through a public repository to support reproducibility and efficient use of produced resources in this study.

We have uploaded these datasets in Researchgate (https://www.researchgate.net/publication/373194650_Trichosanthes_pilosa_datasets Positive_selection_and_relaxed_purifying_selection_contribute_to_rapid_evolution of_male-biased_genes_in_a_dioecious_flowering_plant).

Comments to the authors:

1. I have an issue with the tissue-biased gene expression analysis. Looking at Fig.3, it seems to me there are 3,204 male-biased genes that are expressed at the same level in male buds and mature flowers (same for 5,011 female-biased genes in female buds and flowers), however, only a handful of genes show sex bias between mature male and female flowers. Taking the male-biased genes as an example, if the 3,204 M1BGs experience the same expression levels in mature male flowers and are no longer male-biased when mature male vs female flowers are compared, why there are not found as female tissue biased (F2TGs)? I may be wrong, but one scenario would be that the M1BGs increase their expression in female flowers and become unbiased. However, that increase in expression (low expression in the female buds → higher expression in the female flowers) should classify them as female tissue-biased genes (F2TGs). Can you please clarify how are the M1BGs and F1BGs expressed in the flowers of the opposite sex?

As to Fig. 3A, 3,204 male-biased genes expressed in male floral buds are part of all male-biased genes (3204+286+724=4214), as shown in Fig.2A. However, only 233 male-biased genes (88+1+144=233, Fig.2B and Fig.3B) expressed in male mature flowers. So, they are not expressed at the same level between male floral buds and mature flowers. Only 288 genes are sex-biased (M1BGs), as well as tissue/stage-biased (M1TGs) in male floral buds. M1BGs (4,214 male-biased genes) and F1BGs (5,096 female-biased genes) are 0 overlaps, except for 44,326 unbiasedgenes shown in Fig.2A. That is, F1BGs (5,096 female-biased genes) are low expression or no expression in M1BGs (4,214 male-biased genes). The expression levels of some genes have been shown in Table S14.

1. Paragraph (407-416) describes the analysis of duplicated genes under relaxed selection but there is no mention of this in the results.

In fact, these results have been shown in Table S13. It is not necessary for us to describe them in detail in the results.

1. How did the authors conclude that the identified functions in male flowers make them more adapted to biotic and abiotic environments (line 347-350)? In the paragraph above (line 338-342) the authors describe that female buds are better equipped against herbivores, which are a biotic factor?

Following your concerns, we have revised the manuscript as follows: For line 338-342, we revised the text as “Indeed, functional enrichment analysis in chemical pathways such as terpenoid backbone and diterpenoid biosynthesis indicated that relative to male floral buds, female floral buds had more expressed genes that were equipped to defend against herbivorous insects and pathogens, except for growth and development (Vaughan et al., 2013; Ren et al., 2022) (Fig. S7A and Table S11).” For line 347-350, we revised text as “We also found that male-biased genes with high evolutionary rates in male buds were associated with functions to abiotic stresses and immune responses (Tables S12 and S13), which suggest that male floral buds through rapidly evolving genes are adapted to mountain climate and the environment in Southwest China compared to female floral buds through high gene expression.”

1. Line 417-418: decreasing codon usage bias is linked to decreasing synonymous substitution rates, should this be the opposite?

No. Codon usage bias was positively related to synonymous substitution rates. That is, stronger codon usage bias may be related to higher synonymous substitution rates (Parvathy et al., 2022).

1. Figures and Tables are not standalone and are missing details in the legends. - Fig.2C, which genes are plotted on the heatmap and what is the color scale corresponding to?

All Supplementary figures are missing the descriptions of individual panels (A, B, C,etc.) in the legends. In addition, please add the numbers of observations under boxplots.

Supplementary Fig.5 and 6: Panel B is not a Venn diagram, I suggest removing it from the figures.

Supplementary Fig.7: Should be 'sex-biased genes'. What is the x-axis on the plot?

Supplementary Fig.8: Please add the description of the abbreviations in the legend. - Supplementary Tables S4, S5, S6: Please add information about the foreground and background branches.

Supplementary Table S6, S7, S8, S9, S10: Please add more details about the column headers (what is Model-A, background ω 2a, Unconstrained_1.p, K, which was the foreground branch etc.).

Supplementary Table S11: Please add gene IDs for each KEGG category.

We have revised/fixed these issues following your concerns and suggetions.

Minor comments:

Line 28: 'algae' in place of 'algas'

Line 53-56: Please provide more recent references.

Line65: 'most' instead of 'almost'

Line 86-87: It is not clear from the sentence if the sex-biased expression was detected in flowers compared to leaves, or were the sex-biased genes detected between male and female leaves? Please clarify.

Line 107-108: positive selection is referred to as adaptive evolution, please choose one or the other.

Line 109: 'force' instead of 'forces'

Line 110: 'algae' instead of 'alga'

Line 132: '..mainly distributed from Southwest,' the country is missing.

Line 202: 'protein sequence evolution'?

Line 232: what does the 'number of evolutionary rates' refers to?

Line 253: please provide a reference for the RELAX model.

Line 274: 'relaxed selective male-biased genes' should be 'male-biased genes under relaxed purifying selection'?

Line 318: Please add a sentence explaining why the Cucurbitaceae family is a great model to study the evolution of sexual systems.

Line 321: 'genes' instead of 'gene'.

Line 366: male-biased genes experience 'higher' or 'more rapid' evolutionary rates. line 377: in the present study and in the case of Ectocarpus alga, positive selection plays an important role in male-biased genes evolution, but does not account for the majority of evolutionary change. Therefore, I would not call it a 'primary' force.

Line 477: missing reference for DESeq2 package.

Line 480: 'used'.

Line 498: 'coding sequences'.

Line516: 'to' instead of 'by'.

Line 553: 'the' is repeated twice.

Sorry for the typos and grammatical issues. We have revised them accordingly.

Reviewer #2 (Recommendations For The Authors):

There are two areas for improvement, one empirical and one theoretical.

Empirically, the analyses could be expanded by an attempt to distinguish between genes on the autosomes and the sex chromosomes. Genotypic patterns can be used to provisionally assign transcripts to XY or XX-like behavior when all males are heterozygous and all females are homozygous (fixed X-Y SNPs) and when all females are heterozygous and males are homozygous (lost or silenced Y genes). Comparing such genes to autosomal genes with sex-biased expression would sharpen the results because there are different expectations for the efficacy of selection on sex chromosomes. See this paper (Hough et al. 2014; https://www.pnas.org/doi/abs/10.1073/pnas.1319227111), which should be cited and does in fact identify faster substitution rates in Y-linked genes (and note that pollenexpressed genes, at least, are concentrated on the sex chromosome in this system: https://academic.oup.com/evlett/article/2/4/368/6697528, https://royalsocietypublishing.org/doi/10.1098/rstb.2021.0226).

We have cited Hough et al. 2014 and noticed that several species have been observed to exhibit rapid evolutionary rates of sequences on sex chromosomes compared to autosomes, which has been related to the evolutionary theories of fast-X or fast-Z (lines 482-484).

On the theoretical side, this study is making a very specific intervention, namely identifying more rapid evolutionary rates in genes with male-biased than femalebiased expression in a dioecious plant. The writing in the introduction and the discussion needs to be improved to differentiate between this comparison and similar comparisons, e.g. sex-biased expression in other dioecious plants (76-81), between Xlinked and Y-linked genes (Hough et al. 2014), sex chromosomes and autosome (several studies already cited), gametophytic and sporophytic tissue, and male and female reproductive tissue in hermaphroditic plants. Setting out this distinction early in the introduction will make the specific goals and novelty of this work clearer.

Thank you for your constructive suggestions. We have revised the relevant part of the Introduction accordingly (lines 74-107).

Specific comments by line:

Sorry for the typos or wording issues. We have revised them.

26 - driven not driving

28 - check house style (algae vs algas)

28-29 - consider clarifying the antecedent of "them" (evolutionary forces, not algas) 35 - maybe, but don't the signalling genes involved in stress responses function in many capacities, not just stress? Also, there's evidence that reproductive recognition machinery in plants may ultimately derive from immune function (e.g. https://doi.org/10.1111/j.1469-8137.2008.02403.x), so the GO category "biotic stress" may be too vague

39 - maybe clarify that "for the first time" refers to male rather than female, since there have been other studies in dioecious plants

66-68 - asserting that something is "essential" after describing how rare it is doesn't quite follow, since diecious plants - especially with sex chromosomes - are basically an exception. I agree that understanding the evolution of dioecious plants is important, but this isn't the most compelling way to make that case - perhaps try something else.

137ff - this sentence can be consolidated and streamlined

142 - "floral tissue" rather than "flowers tissue," here and elsewhere

144 - divergence (singular)

235 - "evidence for the contributions of" = "evidences" is unidiomatic 250 - efficiency or efficacy?

300 - why is "inositol" capitalized here and elsewhere?

300ff - are these typical patterns in male tissue in other species?

308 - is that interesting? It seems like exactly what I'd expect. Perhaps start with the unsurprising but reassuring observation (anther and pollen development genes are indeed expressed in male buds) before moving on to the more surprising findings.

319 - remove "the"

321 - genes (plural)

330 - replace "these differences" with "the differences" 336 - perhaps recap proportions / percents here?

340 - unnecessary comma after diterpenoid

341 - this seems like a big leap from the evidence, especially in the absence of supporting information about the chemical defenses of these species and how they differ by sex. Don't terpenoids have a diverse array of functions, not just defense? Here's a review: https://link.springer.com/chapter/10.1007/10_2014_295

We have revised the text as “Indeed, functional enrichment analysis in chemical pathways such as terpenoid backbone and diterpenoid biosynthesis indicated that relative to male floral buds, female floral buds had more expressed genes that were equipped to defend against herbivorous insects and pathogens, except for growth and development (Vaughan et al., 2013; Ren et al., 2022) (Fig. S7A and Table S11)” (lines 373-378).

349 - as mentioned in line 35, this is a big speculative leap. The discussion is the place for speculation, but consider other explanations too. How does the development of flowers work? Are male flowers suppressing or resorbing female primordial organs? Do male flowers in fact senesce faster? perhaps spell out the logic in more detail.

We have revised the text as “In addition, the enrichment in regulation of autophagy pathways could be associated with gamete development and the senescence of male floral buds (Table S14) (Liu and Bassham, 2012; Li et al., 2020; Zhou et al., 2021). In fact, it was observed that male flowers senesced faster (Wu et al., 2011). We also found that homologous genes of two male-biased genes in floral buds (Table S14) that control the raceme inflorescence development (Teo et al., 2014) were highly expressed compared to female floral buds. Taken together, these results indicate that expression changes in sex-biased genes, rather than sex-specific genes play different roles in sexual dimorphic traits in physiology and morphology (Dawson and Geber, 1999).” (lines 390-402).

351 - senescence of, not senescence for

363 - but Hough et al. 2014 did show rapid evolution of Y-linked genes, and those are by definition sex biased ...

391 - perhaps reiterate here that while some sex-BIASED genes did, sex-SPECIFIC genes did not, to avoid confusion

We also revised them accordingly.

Reviewer #3 (Recommendations For The Authors):

1- lines 56-57 : « have facilitated » : this wording confounds correlation with causation. Consider rephrasing as « is associated with »

2- lines 58-60 : vague wording : what are these variations ? e.g. which tissues and stages are generally enriched?

3- line 63 : this sentence is a bit misleading: consider changing it to « Most dioecious plants possess homomorphic sex-chromosomes » [and explain what homomorphic means in this context].

4- line 68 : a reference is missing here. Also perhaps, allude to the fact that sexual selection in plants has long been considered a contentious issue (e.g. https://doi.org/10.1016/j.cub.2010.12.035)

5- lines 72-76 : beyond simply describing the pattern, say what evolutionary processes are revealed by these observations.

6- line 92 : remind the reader what these 5 studies are.

7- line 94-95 : explain why the comparison of vegetative vs vegetative and vegetative vs reproductive tissues is a problem.

The published studies only compared gene expression in vegetative versus vegetative tissues and vegetative versus reproductive tissues. Because it limited our understanding of sexual selection at different floral development stages. Revised accordingly (lines 103-104). We are very interested in flower development stage for sex-biased genes. The datasets could investigate sexual selection using two developmental stage (buds + mature flowers).

8- line 100 « Evolutionary dynamic analyses » : this wording is vague

9- line 110 : brown algae are NOT plants

10- line 137-140 or in M&M : needs to describe somewhere how the male flowers differ from the female flowers and vice-versa: are the whole morphological structures related to female (male) reproduction entirely missing, or is their development arrested later on and they are still present but simply not producing gametes? This has consequences for the interpretation of the genes they express.

We have revised the typos or wording issues accordingly. However, because the sampled floral buds were equal or less than 3 mm in size, we did not observe much morphological structural difference. Indeed, the male and female flowers at antheses were markedly different in this dioecious plant as shown in Fig. 1. Additionally, we found that dioecy is the ancestral state of Trichosanthes, and transitions to monoecy (Guo et al., 2020) based on our analysis (not shown in this study), which suggest that in the early stages of flower development, female floral buds may tend to masculinize, and vice versa (Fig. 2C).

11- line 152 : it is important to be very transparent on the sample sizes here: « from three females and three males of the dioecious ... »

12- line 153 : along the same line, explain here why a de novo transcriptome had to be generated here: « In the absence of an assembled reference genome for this nonmodel species, we de novo assembled ... »

13- line 164-165 : « we have generated high-quality reference trancriptomes » : I am not entirely convinced of the quality of the transcriptome obtained without a reference genome, so I suggest simply removing this subjective sentence.

Our assessment suggested that we have generated high-quality reference transcriptomes in the absence of a reference genome, which will be the next step of our work.

14- line 169 : briefly explain the criteria used to call differentially expressed genes. Given the threshold (log-fold change >=1.3 if I read the figure correctly, but the M&M says >=1), explain how it was chosen.

Sorry, you may have misunderstood the X, Y coordinates. The value of y coordinate represents -log10(FDR), and the value of x coordinate represents log2 (Fold Change).

15- line 174 : Not clear to me how Fig2C is « revealing strong sexual dimorphism », since genes cluster neither by sex nor by tissue. This should be explained more clearly.

16- line 174-177 : The fact that more ex-biased genes were identified in early buds than in mature flowers is an interesting observation that could be given more prominence in the manuscript, but it is not really explained. Could it reflect the fact that more genes are expressed in early buds because meiotic processes happen early in flower development? Also, the genes involved in male and female organ cell fate determination might also be expected to be expressed early, with mostly organ growth genes being expressed in the mature flower.

17- line 181 : a wrap-up sentence might be useful here to drive the point home that sex-bias is more prevalent in buds than mature flowers.

18- line 184 : « tissue-biased » : a more appropriate wording here would be « stagebiased », no ? These are indeed the same tissues but at different developmental stages.

19- line 183-195 : this section could benefit from setting clear expectations in a hypothesis testing framework laying out the reasons to expect a different bias between stages and between sexes. How does that fit with the level of morphological divergence between sexes (relates to my point 10 above).

20- line 197-204. A number of essential pieces of information are missing here: how many species, how divergent, say that one other is dioecious, and precise their relative phylogenetic placement (which is important to understand the models used below). Maybe adding a phylogeny of these species in Figure 4 could be useful. Also, briefly explain the « two-ratio » and « free-ratio » models here.

21- line 196 and following: In these analyses, I could not understand the rationale for keeping buds vs mature flowers as separate analyses throughout. Why not combine both and use the full set of genes showing sex-bias in any tissue? This would increase the power and make the presentation of the results a lot more straightforward.

As you pointed earlier (in the public review, paragraphy 2), “the dataset is interesting in that it focuses on reproductive rather than somatic tissues (which makes sense to investigate sexual selection), and includes more than a single developmental stage (buds + mature flowers)”, we totally agree with your points and were very interested in floral development stages for sex-biased genes.

22- line 216 : say explicitly that the reason for not detecting a significant difference in spite of a relatively large effect size is probably related to the low number of genes, conferring low statistical power to detect a difference. An important feature also not highlighted here is that the trend (though not significant) is in the opposite direction than in the buds, and that both the 2-ratio and the free-ratio models agree on these inverted trends. This could be the basis for an interesting comparison.

Thank you for your suggestions.

23- line 220 : needs to explain more clearly how this « free-ratio » differs from the « two-ratio » model.

24- line 232-234 : I don't see why this is necessary. Why not combine both? See also my comment 21 above.

25- line 237 : the «A-model » was not defined before.

26- line 237 : « male-biased » is missing after 343.

27- line 253-258 : briefly explain what these other models are based on and how they are not redundant and instead complement the previous analyses and each other. 28- line 266-268 : the use of a more precise set of codons for male-biased genes than the others (if I understood correctly) could also be interpreted as another sign of stronger selective constraint, no?

Codon usage bias is influenced by many factors, such as levels of gene expression. Highly expressed genes have a stronger codon usage bias and could be encoded by optimal codons for more efficient translation (Frumkin et al., 2018; Parvathy et al., 2022).

29- line 269-291 : removing genes on a post-hoc basis seems statistically suspicious to me. I don't think your analysis has enough power to hand-pick specific categories of genes, and it is not clear what this brings here. I suggest simply removing these analyses and paragraphs.

30- line 325 : say whether this patterns parallels / or not those in animals.

31- line 335 : yes, these biological pieces of information are important and should be given in the introduction already.

32- the discussion should focus at some point on the observation that more femalebiased genes are found in general, but that male-biased genes seem to be under greater selection. How do you reconcile these two apparently contradictory observations?

We found that male-biased genes with high evolutionary rates in male floral buds were associated with functions to abiotic stresses and immune responses (Tables S12 and S13), which suggests that male floral buds through rapidly evolving genes are adapted to mountain climate and the environment in Southwest China compared to female floral buds through high gene expression (lines 387-390).

33- line 355 : not clear how this follows from the previous sentences.

34- line 356-358 : vagiue. not clear what the message of this sentence is.

35- line 378-383 : say that these conclusions rely on the quality of gene annotation in this non-model species, which is probably pretty low (just like any other non-model species).

36- line 403 : this conclusion seems far-fetched. Explain how exactly you reached this conclusion.

37- line 406-416: these speculations on the role of paralogs seem unnecessary, in particular since the de novo transcriptome onto which all analyses are based cannot distinguish orthologs from paralogs.

38- line 417-424. The discussion should not contain new results.

39- line 510 : why were genes with dN/dS >2 discarded here? This might strongly bias the comparison, no? This needs to be properly justified.

40- lines 516-523 : references to the models are missing.

41- line 527: « omega = 1.5 » : why/how was this arbitrary threshold chosen?

42- Fig 2 : write out « buds » and « mature flowers » on top of the graphs

43- Fig 4 : add a phylogeny of the species showing the branch being compared. Also, add the number of genes in each category on each plot.

Thanks, we revised/fixed these issues accordingly.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Sex-biased genes offer insights into the evolution of sexual dimorphism. Sex-biased genes, especially those with male bias, show elevated evolutionary rates of protein sequences driving by positive selection and relaxed purifying selection in animals. Although rapid sequence evolution of sex-biased genes and evolutionary forces have been investigated in animals and brown algas, less is known about them in dioecious angiosperms. In this study, we separately compared the expression of sex-biased gene between female and male flower buds and between female and male mature flowers in dioecious Trichosanthes pilosa (Cucurbitaceae). In floral buds, sex-biased gene expression was pervasive, and had significantly different roles on sexual dimorphism such as physiology. We observed higher rates of sequence evolution for male-biased genes in floral buds compared to female-biased and unbiased genes. Male-biased genes under positive selection were mainly related to abiotic and biotic stress, suggesting that high evolutionary rates are driven by adaptive evolution. Additionally, relaxed purifying selection may contribute to accelerated evolution in male-biased genes generated by gene duplication. Our findings, for the first time in dioecious angiosperms, suggest evident rapid evolution of male-biased genes, advance our understanding of the patterns and forces driving the evolution of sexual dimorphism in dioecious plants.

Revised version: Sex-biased genes offer insights into the evolution of sexual dimorphism. Sex-biased genes, especially those with male bias, show elevated evolutionary rates of protein sequences driven by positive selection and relaxed purifying selection in animals. Although rapid sequence evolution of sex-biased genes and evolutionary forces have been investigated in animals and brown algae, less is known about evolutionary forces in dioecious angiosperms. In this study, we separately compared the expression of sex-biased genes between female and male floral buds and between female and male flowers at anthesis in dioecious Trichosanthes pilosa (Cucurbitaceae). In floral buds, sex-biased gene expression was pervasive, and had significantly different roles in sexual dimorphism such as physiology. We observed higher rates of sequence evolution for male-biased genes in floral buds compared to female-biased and unbiased genes. Male-biased genes under positive selection were mainly associated with functions to abiotic stress and immune responses, suggesting that high evolutionary rates are driven by adaptive evolution. Additionally, relaxed purifying selection may contribute to accelerated evolution in male-biased genes generated by gene duplication. Our findings, for the first time in angiosperms, suggest evident rapid evolution of male-biased genes, advance our understanding of the patterns and forces driving the evolution of sexual dimorphism in dioecious plants.

Introduction

Earlier version: Sexual dimorphism presents different morphological, ecological and physiological traits in gonochoristic animals and dioecious plants, despite male and female individuals sharing the same genome except for sex chromosomes (or sex-determining loci) (Mank, 2009; Barrett and Hough, 2013). Such sexual dimorphisms usually arise from differential expression of genes in both sexes, i.e., sex-biased genes (including sex-specific genes exclusively in one sex) that are located on autosomal chromosomes and sex chromosomes/or sex-determining regions (Ellegren and Parsch, 2007; Parsch and Ellegren, 2013; Charlesworth, 2018). Recently, some studies have begun to explore the strength and impact of evolutionary forces that shape different sexually dimorphic traits through sex-biased gene expression (Mank, 2009; Rowe et al., 2018). Previous studies revealed that sex-biased gene expressions have facilitated the evolution of sexual dimorphisms in many animals. However, the extent to which this occurs varies greatly among taxa, tissues, and developmental stages (Mank, 2017; Khodursky et al., 2020; Hsu et al., 2020; Lichilin et al., 2021; Toubiana et al., 2021; Djordjevic et al., 2022). Unlike most animals, the vast majority (~90%) of angiosperms are hermaphroditic plants, while only a small fraction (~5%) are dioecious plants in which individuals have exclusively male or female reproductive organs (Renner, 2014). Most dioecious plants do not possess heteromorphic sex chromosomes. Furthermore, sexual dimorphism in dioecious plants is less common and less conspicuous than that in almost animals (Barrett and Hough, 2013). Hence, the study of sex-biased gene expression is of great interest to plant evolutionary biologists, as it is essential to understand the evolution of sexual dimorphism in dioecious plants.

Revised version: Sexual dimorphism presents different morphological, ecological and physiological traits in gonochoristic animals and dioecious plants, despite male and female individuals sharing the same genome except for sex chromosomes (or sex-determining loci) (Mank, 2009; Barrett and Hough, 2013). Such sexual dimorphisms usually arise from differential expression of genes in both sexes, i.e., sex-biased genes (including sex-specific genes exclusively in one sex) that are located on autosomal chromosomes and sex chromosomes/or sex-determining regions (Ellegren and Parsch, 2007; Parsch and Ellegren, 2013; Charlesworth, 2018b). Recently, some studies have begun to explore the strength and impact of evolutionary forces that shape different sexually dimorphic traits through sex-biased gene expression (Mank, 2009; Rowe et al., 2018; Naqvi et al., 2019; Tosto et al., 2023). Previous studies revealed that sex-biased gene expressions were associated with the evolution of sexual dimorphisms in many animals, although the extent to which exhibits great variation among taxa, tissues, and development stages (Mank, 2017; Hsu et al., 2020; Khodursky et al., 2020; Lichilin et al., 2021; Toubiana et al., 2021; Djordjevic et al., 2022; Yue et al., 2023). Unlike most animals, the vast majority (∼90%) of flowering plants (angiosperms) are hermaphroditic, while only a small fraction (∼5%) are dioecious in which individuals have exclusively male or female reproductive organs (Renner, 2014). Most dioecious plants possess homomorphic sex-chromosomes that are roughly similar in size when viewed by light microscopy. Furthermore, sexual dimorphism in dioecious plants is less common and less conspicuous than that in most animals (Barrett and Hough, 2013). Hence, the study of sex-biased gene expression is of great interest to plant evolutionary biologists, as it is necessary to understand the evolution of sexual dimorphism in dioecious plants (Moore and Pannell, 2011).

Earlier version: A common pattern that has emerged from previous studies is that sex-biased genes, particularly male-biased genes, tend to evolve rapidly in protein sequence (the ratio of non-synonymous to synonymous substitutions, dN/dS) compared to unbiased genes (Ellegren and Parsch, 2007; Grath and Parsch, 2016). The rapid evolution of male-biased genes was first observed in Drosophila melanogaster (Zhang et al., 2004; Zhang and Parsch, 2005) and has been supported by recent investigations in a wider range of animals (Pro-schel et al., 2006; Mank et al., 2007; Mank, 2017; Papa et al., 2017; Catalan et al., 2018; Toubiana et al., 2021). In recent years, there have been growing studies on the expression dynamics and molecular evolutionary rates of sex-biased genes in flowering plants, including Silene latifolia (Zemp et al., 2016), Salix viminalis (Darolti et al., 2018), Mercurialis annua (Cossard et al., 2019), Populus balsamifera (Sanderson et al., 2019), and Leucadendron plants (Scharmann et al., 2021). However, despite such advances, the molecular evolution pattern of sex-biased genes in plants remains inconsistent among the studied plant species (Muyle, 2019; Veltsos, 2019). In dioecious plants such as Mercurialis annua and Leucadendron, Cossard et al., (2019) and Scharmann et al., (2021) found no significant differences in evolutionary rates of proteins among female-biased, male-biased and unbiased genes, although the expression of sex-biased genes was highly different from unbiased genes in leaves. Similar patterns have also been reported in dioecious Populus balsamifera, where evolutionary rates of male-biased, female-biased and unbiased genes did not differ in reproductive tissues (Sanderson et al., 2019). However, in the dioecious plant Salix viminalis, male-biased genes have significantly lower evolutionary rates of proteins than female-biased and unbiased genes in catkin tissues (Darolti et al., 2018). To our knowledge, only five studies have investigated the expression differences and protein evolutionary rates of sex-biased genes in dioecious angiosperms. Moreover, these studies only compared gene expression in vegetative vs. vegetative tissues and vegetative vs. reproductive tissues. Therefore, more studies and taxa are needed to explore the common patterns of sequence evolution in sex-biased genes, with more focus on comparing sex-biased gene expression in reproductive vs. reproductive tissues, e.g., different floral development stages in dioecious angiosperms.

Revised version: A common pattern that has emerged from previous studies is that sex-biased genes, particularly male-biased genes, tend to evolve rapidly in protein sequence (the ratio of non-synonymous to synonymous substitutions, dN/dS) compared to unbiased genes (Ellegren and Parsch, 2007; Grath and Parsch, 2016). The rapid evolution of male-biased genes was first observed in Drosophila melanogaster (Zhang et al., 2004; Zhang and Parsch, 2005) and has been supported by recent investigations in a wider range of animals (Pröschel et al., 2006; Mank et al., 2007; Mank, 2017; Papa et al., 2017; Catalan et al., 2018; Toubiana et al., 2021). In recent years, there have been growing studies on the expression dynamics and molecular evolutionary rates of sex-biased genes in flowering plants, including hermaphroditic Arabidopsis thaliana (Gossmann et al., 2014; 2016), Solanum (Moyle et al., 2021), and dioecious Silene latifolia (Zemp et al., 2016), Salix viminalis (Darolti et al., 2018), Mercurialis annua (Cossard et al., 2019), Populus balsamifera (Sanderson et al., 2019), and Leucadendron (Scharmann et al., 2021). However, despite such advances, the molecular evolution pattern of sex-biased genes in plants remains inconsistent among the studied plant species (Muyle, 2019; Veltsos, 2019). In dioecious plants such as Mercurialis annua and Leucadendron, Cossard et al., (2019) and Scharmann et al., (2021) found no significant differences in evolutionary rates of proteins among female-biased, male-biased and unbiased genes detected between male and female plants leaf tissues, although the expression of sex-biased genes was highly different from unbiased genes in leaves. Similar patterns have also been reported in dioecious Populus balsamifera, where evolutionary rates of male-biased, female-biased and unbiased genes did not differ in reproductive tissues (Sanderson et al., 2019). However, in the dioecious Salix viminalis, male-biased genes have significantly lower evolutionary rates of proteins than female-biased and unbiased genes in catkin tissues (Darolti et al., 2018). To our knowledge, only above mentioned five studies have investigated the expression differences and protein evolutionary rates of sex-biased genes in dioecious angiosperms. Moreover, these studies only compared gene expression in vegetative versus vegetative tissues and vegetative versus reproductive tissues, limiting our understanding of sexual selection at different floral development stages. Therefore, more studies and taxa are needed to explore the common patterns of sequence evolution in sex-biased genes, with more focus on comparing sex-biased gene expression in reproductive versus reproductive tissues, e.g., different floral development stages in dioecious angiosperms.

Earlier version: Evolutionary dynamic analyses indicate that different evolutionary forces impact the rate of sequence evolution of sex-biased genes. These forces include positive selection, which promotes the spread and adaptive fixation of beneficial alleles; sexual selection, which results from male-male competition or female choice; and relaxed purifying selection, which reduces the removal of deleterious mutations (Grath and Parsch, 2016; Mank, 2017; Dapper and Wade, 2020). For example, in animal systems, particularly in Drosophia, the elevated sequence divergence rates of male-biased genes have often been interpreted as the signature of positive selection due to adaptive evolution, suggesting that sexual selection is the primary evolutionary forces (Pro-schel et al., 2006; Assis et al., 2012). However, studies in plants have rarely reported elevated rates of sex-biased genes, except for brown alga, in which female-biased and/or male-biased genes exhibiting higher evolutionary rates than unbiased genes suggests that rapid evolution is partly driven by adaptive evolution or sexual selection (Lipinska et al., 2015; Cossard et al., 2022; Hatchett et al., 2023).

Revised version: Evolutionary analyses indicate that different driving forces impact the rate of sequence evolution of sex-biased genes. These forces include positive selection, which promotes the spread and adaptive fixation of beneficial alleles; sexual selection, which results from male-male competition or female choice; and relaxed purifying selection, which reduces the removal of deleterious mutations (Grath and Parsch, 2016; Mank, 2017; Dapper and Wade, 2020). For example, in animal systems, particularly in Drosophia, the elevated sequence divergence rates of male-biased genes have often been interpreted as the signature of adaptive evolution, suggesting that sexual selection is the primary evolutionary force (Pröschel et al., 2006; Assis et al., 2012). In brown algae, female-biased and/or male-biased genes exhibiting higher evolutionary rates than unbiased genes suggest that rapid evolution is partly driven by adaptive evolution or sexual selection (Lipinska et al., 2015; Cossard et al., 2022; Hatchett et al., 2023). However, studies in plants have rarely reported elevated rates of sex-biased genes.

An alternative explanation for the rapid evolution of sex-biased genes is a relaxation of purifying selection due to reduced constraints (Lahti et al., 2009; Dapper and Wade, 2020). In the model plant Arabidopsis thaliana, pollen genes were found to be evolving faster than sporophyte-specific genes due to relaxed purifying selection associated with the transition from outcrossing to selfing (Harrison et al., 2019). These trends were recently confirmed in Arabis alpina, which exhibits mating system variation across its distribution, suggesting that the efficacy of purifying selection on male gametophyte-expressed genes was significantly weaker in inbred populations (Gutierrez-Valencia et al., 2022). Together, these findings in plants reinforce the idea that both adaptive (e.g., positive selection, sexual selection) and non-adaptive (e.g., relaxed selection) evolutionary processes differentially impact the sequence evolution of sex-biased genes. Hence, investigating the potential contribution of selection forces to the emergence of specific evolutionary patterns of sex-biased genes within a focal species is of great interest.

Earlier version: In the family Cucurbitaceae, there are about 96 genera and 1,000 species, about 50% of species are dioecious, and 50% are monoecious (Schaefer and Renner, 2011). Phylogenetic analyses of Cucurbitaceae suggest that dioecy is the ancestral state of the family, but transitions frequently to monoecy (Zhang et al., 2006). Trichosanthes pilosa (2n = 22, Cucurbitaceae) is mainly distributed from Southwest, Southeast China to Japan, New Guinea and Western Australia. It was suggested to have originated in the late Miocene (c.13 Ma) (de Boer et al., 2012; 2015). Trichosanthes pilosa (synonym: T. ovigera) is a perennial dioecious vine that reproduces sexually and possesses a pair of heteromorphic sex chromosomes XX/XY (Ming et al., 2011). The species exhibits strong sexual dimorphism in the floral morphological structures and the life-history traits, such as racemose male flower and solitary female flower (Fig. 1), early-flowering, caducous male flower and late-flowering, long-lived female flower (Wu et al., 2011).

Revised version: In the family Cucurbitaceae, there are about 96 genera and 1,000 species, about 50% of species are dioecious, and the others are monoecious (Schaefer and Renner, 2011). Phylogenetic analyses of Cucurbitaceae suggest that dioecy is the ancestral state of the family, but transitions frequently to monoecy (Zhang et al., 2006). Trichosanthes pilosa (synonym: T. ovigera, 2n = 22, Cucurbitaceae) is mainly distributed from Southwest and Southeast China to Japan, extending to Southeast Asia, New Guinea and Western Australia. It was suggested to have originated in the late Miocene (ca. 8-6 million-year ago) (de Boer et al., 2012; 2015; Guo et al., 2020). Trichosanthes pilosa is a perennial, night-flowering, insect-pollinated dioecious vine that reproduces sexually and possesses a pair of heteromorphic sex chromosomes XX/XY (Ming et al., 2011). It’s male and female flowers exhibit strong sexual dimorphism in floral morphological and phenological traits, such as racemose versus solitary (Fig. 1), early-flowering versus late-flowering, and caducous versus long-lived (Wu et al., 2011).

Earlier version: Floral buds and mature flowers of females (A, B) and males (C) in dioecious Trichosanthes pilosa.

Revised version: Floral buds and flowers at anthesis of females (A, B) and males (C) in Trichosanthes pilosa.

Earlier version: To understand the evolution of sex-biased genes in dioecious Trichosanthes pilosa, we collected floral buds and mature flowers tissues from male and female individuals and characterized their expression profiles using Illumina RNA sequencing. Our primary objectives are to 1) compare expression divergences between males and females at two floral developmental stages; 2) explore whether there are differences in the evolutionary rates of proteins among female-biased, male-biased and unbiased genes; and if so, 3) determine the main selective forces that contribute to the differentiation of sequence evolution rates among gene categories.

Revised version: To understand the evolution of sex-biased genes in dioecious T. pilosa, we collected floral buds and flowers at anthesis from male and female individuals and characterized their expression profiles using Illumina RNA sequencing. Our primary objectives are to 1) compare expression divergence between males and females at two floral development stages; 2) explore whether there are differences in the evolutionary rates of proteins among female-biased, male-biased and unbiased genes; and if so, 3) determine the main selective forces that contribute to the differentiation of sequence evolution rates among gene categories.

Results

Transcriptome sequencing, de novo assembly and annotation

Earlier version: Using whole transcriptome shotgun sequencing, we sequenced 12 floral buds and mature flower tissues from females and males of dioecious T. pilosa and generated a total of nearly 276 million clean reads (Table S1). We then performed de novo assembly of transcripts from all the clean reads, followed by clustering and filtering analysis, resulting in 59,051 unigenes (Fig. S1A). To evaluate the quality of the assembled unigenes, we compared them against protein databases such as NR, KEGG, Swissport, PFAM, and GO using BLASTP and nucleotide database NT using BLASTN (Table S2). The e-value distribution of the best hits in NR database suggested that 47,241 unigenes (80%) had strong homology, with an e-value smaller than 1.0e-15 (Fig. S1B). The majority of unigenes were annotated by homologs in other species of Cucurbitaceae, such as Momordica charantia, Cucumis melo, Cucurbita pepo, Cucurbita moschata and Cucurbita maxima (Fig. S1C). Furthermore, we used BUSCO assessments based on embryophyta_odb9 database, which showed the completeness of the reference transcriptome at 78.8% (Table S2). Overall, our analyses suggested that we have generated high-quality reference transcriptomes.

Revised version: Using whole transcriptome shotgun sequencing, we sequenced 12 floral buds and flowers at anthesis from three females and three males of dioecious T. pilosa and generated a total of nearly 276 million clean reads (Table S1). Due to the absence of a reference genome, we performed de novo assembly of transcripts from all the clean reads, followed by clustering and filtering analysis, resulting in 59,051 unigenes (Fig. S1A). To evaluate the quality of the assembled unigenes, we used BUSCO assessments based on embryophyta_odb10 database, which showed the completeness of the reference transcriptome at 89.7% (Table S2). We then annotated them against protein databases including NR, KEGG, Swissport, PFAM, and GO using BLASTP and nucleotide database NT using BLASTN (Table S2). The e-value distribution of the best hits in NR database suggested that 47,241 unigenes (80%) had strong homology, with an e-value smaller than 1.0e-15 (Fig. S1B). The majority of unigenes were annotated by homologs in species of Cucurbitaceae, such as Momordica charantia (16.3%), Cucumis melo (11.9%), Cucurbita pepo (11.9%), Cucurbita moschata (11.5%) and Cucurbita maxima (10.1%) (Fig. S1C). Overall, our assessment suggested that we have generated high-quality reference transcriptomes.

Expression characteristics of sex-biased genes

Earlier version: We mapped the RNA-Seq reads of floral buds and mature flowers onto the reference transcriptome in dioecious T. pilosa, which resulted in approximately 75% read mappings per sample (Table S3). In floral buds, we identified 5,096 (9.50%) female-biased genes and 4,214 (7.86%) male-biased genes (Fig. 2A). In contrast, only 380 (0.70%) female-biased genes and 233 (0.43%) male-biased genes were detected in mature flowers (Fig. 2B). Using hierarchical clustering analysis, we evaluated different levels of gene expression across sexes and tissues, revealing strong sexual dimorphism (Fig. 2C). Furthermore, we observed that the number of sex-biased genes in floral buds was approximately 15 times higher than in mature flowers, indicating that sex-biased genes associated with sex differentiation and sexually dimorphic traits are predominantly expressed in floral buds. We also analyzed sex-specific genes that were exclusively expressed in floral buds and mature flowers of one sex. In floral buds, we found 253 out of 5,096 (4.96%) female-specific genes and 465 out of 4,214 (11.03%) male-specific genes. However, in mature flowers, we only identified 26 out of 380 (6.84%) female-specific genes and 52 out 233 (22.32%) male-specific genes.

Revised version: We mapped the RNA-Seq reads of floral buds and flowers at anthesis onto the reference transcriptome in dioecious T. pilosa, which resulted in approximately 75% read mappings per sample (Table S3). In floral buds, we identified 5,096 (9.50%) female-biased genes and 4,214 (7.86%) male-biased genes (Fig. 2A). In contrast, only 380 (0.70%) female-biased genes and 233 (0.43%) male-biased genes were detected in flowers at anthesis (Fig. 2B). Using hierarchical clustering analysis, we evaluated different levels of gene expression across sexes and tissues (Fig. 2C), revealing that female floral buds maybe tend to masculinize, and vice versa in the early stages of floral development. Furthermore, we observed that the number of sex-biased genes in floral buds was approximately 15 times higher than in flowers at anthesis, indicating that sex-biased genes associated with meiotic processes, sex differentiation and sexually dimorphic traits are predominantly expressed in floral buds. We also analyzed sex-specific genes that were exclusively expressed in floral buds and flowers at anthesis of one sex. In floral buds, we found 253 out of 5,096 (4.96%) female-specific genes and 465 out of 4,214 (11.03%) male-specific genes. However, in flowers at anthesis, we only identified 26 out of 380 (6.84%) female-specific genes and 52 out of 233 (22.32%) male-specific genes. Taken together, sex bias is more prevalent in floral buds than in flowers at anthesis.

Earlier version: Sex-biased gene expression for floral buds and mature flowers in males and females of T. pilosa. Volcano plots of average expression between female-biased, male-biased and unbiased genes in floral buds (A) and mature flowers (B). Heatmap of sex-biased gene expression (C) using hierarchical clustering analysis.

Revised version: Sex-biased gene expression for floral buds and flowers at anthesis in males and females of Trichosanthes pilosa. Volcano plots of average expression between female-biased, male-biased and unbiased genes in floral buds (A) and flowers at anthesis (B). M1 and F1 indicate male and female floral buds; M2 and F2 indicate male and female flowers at anthesis. The value of y coordinate represents -log10(FDR), and the value of x coordinate represents log2(Fold Change) identified by DESeq2. Heatmap of sex-biased gene expression (C) using hierarchical clustering analysis. Hierarchical gene clustering is based on Euclidean distance with an average of log2(FPKM) for differentially expressed genes. The color gradient represents from high to low (from red to green) gene expression.

Earlier version: Tissue-biased gene expression

Revised version: Tissue-biased/stage-biased gene expression

Earlier version: We compared the expression levels of transcripts in floral bud tissues and mature flower tissues within each sex to identify genes with tissue-biased expression. In male plant, the number of tissue-biased genes in mature flowers was 1,040 higher than that in floral buds (Fig. 3A and 3B). However, in female plant, the number of tissue-biased genes in mature flowers was only 536 more than that in floral buds (Fig. 3C and 3D). Our results indicated that males had a higher tissue-bias relative to females. We also identified sex-biased genes that were expressed in both types of tissues by comparing tissue-biased genes with male-biased and female-biased genes, respectively. Few of female-biased genes overlapped with tissue-biased genes, accounting for only 85 out of 5,096 (1.67%) in floral buds (Fig. 3C) and 5 out of 380 (1.32%) in mature flowers (Fig. 3D). However, a significant proportion of male-biased genes overlapped with tissue-biased genes, with 1,010 out of 4,214 (23.97%) in floral buds (Fig. 3A) and 145 out of 233 (62.23%) in mature flowers (Fig. 3B).

Revised version: We compared the expression levels of transcripts in floral buds and flowers at anthesis within each sex to identify genes with tissue-biased expression. In male plants, the number of tissue-biased genes in flowers at anthesis was 1,040 higher than that in floral buds (Fig. 3A and 3B). However, in female plants, the number of tissue-biased genes in flowers at anthesis was only 536 more than that in floral buds (Fig. 3C and 3D). Our results indicated that males had a higher tissue-bias relative to females. We also identified sex-biased genes that were expressed in both types of tissues by comparing tissue-biased genes with male-biased and female-biased genes, respectively. Few female-biased genes overlapped with tissue-biased genes, accounting for only 85 out of 5,096 (1.67%) in floral buds (Fig. 3C) and 5 out of 380 (1.32%) in flowers at anthesis (Fig. 3D). However, a significant proportion of male-biased genes overlapped with tissue-biased genes, with 1,010 out of 4,214 (23.97%) in floral buds (Fig. 3A) and 145 out of 233 (62.23%) in flowers at anthesis (Fig. 3B).

Earlier version: The overlap between sex-biased and tissue-biased genes in two type sexes and tissues. Male-biased genes in floral buds (M1BGs) (A) or mature flowers (M2BGs) (B) overlapped with tissue-biased genes of floral buds (M1TGs) and mature flowers (M2TGs). Female-biased genes in floral buds (F1BGs) (C) or mature flowers (F2BGs) (D) overlapped with tissue-biased genes of floral buds (F1TGs) and mature flowers (F2TGs).

Revised version: The overlap between sex-biased and tissue-biased genes in two types of sexes and tissues. Male-biased genes in floral buds (M1BGs) (A) or flowers at anthesis (M2BGs) (B) overlapped with tissue-biased genes of floral buds (M1TGs) and flowers at anthesis (M2TGs). Female-biased genes in floral buds (F1BGs) (C) or flowers at anthesis (F2BGs) (D) overlapped with tissue-biased genes of floral buds (F1TGs) and flowers at anthesis (F2TGs).

Elevated protein evolutionary rates of male-biased genes in floral buds

Earlier version: We identified 1,145 female-biased, 343 male-biased, and 2,378 unbiased one-to-one orthologous groups (OGs) from floral buds in three additional species of Trichanthes, T. anguina, T. pilosa, T. kirilowii, together with Luffa cylindrica. Additionally, we detected 45 female-biased, 13 male-biased, and 3,782 unbiased one-to-one OGs from mature flowers in all four species. To quantify the rates of protein sequences, we separately calculated ω values for each sex-biased and unbiased orthologous gene using two branch models in floral buds and mature flowers (Figs. 4 and S2).

Revised version: In four species with phylogenetic relationships (((T. anguina, T. pilosa), T. kirilowii), Luffa cylindrica), including dioecious T. pilosa, dioecious T. kirilowii, monoecious T. anguina in Trichosanthes, together with monoecious Luffa cylindrica, we identified 1,145 female-biased, 343 male-biased, and 2,378 unbiased one-to-one orthologous groups (OGs) from floral buds. Additionally, we detected 45 female-biased, 13 male-biased, and 3,782 unbiased one-to-one OGs from mature flowers in all four species. To quantify the rates of protein sequences, we separately calculated ω values for each sex-biased and unbiased orthologous gene using ‘two-ratio’ and ‘free-ratio’ branch models in juvenile and mature flowers (Figs. 4 and S2).

Earlier version: Violin plots of dN/dS values (0< ω < 2) of female-biased, male-biased and unbiased genes in floral buds and mature flowers of the dioecious T. pilosa. White dot indicates the median of dN/dS values for sex-biased and unbiased genes. Wilcoxon rank sum tests are used to test for significant differences (***P < 0.0005, **P < 0.005 and *P < 0.05). The distributions of dN/dS values for female-biased, male-biased and unbiased genes in floral buds (A) and mature flowers (B) using ‘two-ratio’ branch model. The distributions of dN/dS values for female-biased, male-biased and unbiased genes in floral buds (C) and mature flowers (D) using ‘free-ratio’ branch model.

Revised version: Violin plots of dN/dS values (0 < ω < 2) of female-biased, male-biased and unbiased genes in floral buds and flowers at anthesis of Trichosanthes pilosa. White dot indicates the median of dN/dS values for sex-biased and unbiased genes. Wilcoxon rank sum tests are used to test for significant differences (***P < 0.0005, **P < 0.005 and *P < 0.05). The distributions of dN/dS values for female-biased, male-biased and unbiased genes in floral buds (A) and flowers at anthesis (B) using ‘two-ratio’ branch model. The distributions of dN/dS values for female-biased, male-biased and unbiased genes in floral buds (C) and flowers at anthesis (D) using ‘free-ratio’ branch model.

Earlier version: In the results of ‘two-ratio’ branch model, the median of ω values in female-biased, male-biased and unbiased genes were 0.227, 0.257 and 0.230 in floral buds, respectively (Fig. 4A and Table S4). We observed that male-biased genes had a 13.22% and 11.74% higher median than female-biased and unbiased genes in floral buds, respectively. The difference in the distribution of ω values between female-biased versus male-biased genes (P = 0.0021) and male-biased versus unbiased genes (P = 0.0051) was statistically significant in Wilcoxon rank sum tests. However, we did not find a significant difference in ω values between female-biased and unbiased genes in floral buds (Wilcoxon rank sum test, P = 0.4618). In mature flowers, the median of ω values for female-biased, male-biased, and unbiased genes were 0.269, 0.177 and 0.231, respectively (Fig. 4B and Table S4). However, we found that there was no statistically significant difference in the distribution of ω values using Wilcoxon rank sum tests for female-biased versus male-biased genes (P = 0.0556), female-biased versus unbiased genes (P = 0.0796), and male-biased versus unbiased genes (P = 0.3296).

Revised version: The two-ratio branch model, where the foreground (dioecious branches) has a different ω value relative to the background (all other branches), is better supported than the fixed-ratio branch model, where the same ω value was in all branches. In the results of ‘two-ratio’ branch model, the median of ω values in female-biased, male-biased and unbiased genes were 0.227, 0.257 and 0.230 in floral buds, respectively (Fig. 4A and Table S4). We observed that male-biased genes had a 13.22% and 11.74% higher median than female-biased and unbiased genes in floral buds, respectively. The difference in the distribution of ω values between female-biased versus male-biased genes (P = 0.0021) and male-biased versus unbiased genes (P = 0.0051) was statistically significant in Wilcoxon rank sum tests. However, we did not find a significant difference in ω values between female-biased and unbiased genes in floral buds (Wilcoxon rank sum test, P = 0.4618). In flowers at anthesis, the median of ω values for female-biased, male-biased, and unbiased genes were 0.269, 0.177 and 0.231, respectively (Fig. 4B and Table S4). However, we found that there was no statistically significant difference in the distribution of ω values using Wilcoxon rank sum tests for female-biased versus male-biased genes (P = 0.0556), female-biased versus unbiased genes (P = 0.0796), and male-biased versus unbiased genes (P = 0.3296).

Earlier version: The ‘free-ratio’ branch model yielded interesting results. In floral buds, the median ω values for female-biased, male-biased, and unbiased genes were 0.222, 0.265 and 0.226, respectively (Fig. 4C and Table S5). Male-biased genes had a significantly higher median relative to female-biased genes (19.37% higher, Wilcoxon rank sum test, P = 0.0009) and unbiased genes (17.26% higher, Wilcoxon rank sum test, P = 0.0004) in floral buds. However, there was no significant difference in ω values between female-biased and unbiased genes (Wilcoxon rank sum test, P = 0.9862). In mature flowers, the median ω values for female-biased, male-biased, and unbiased genes were 0.300, 0.148 and 0.227, respectively (Fig. 4D and Table S5). Female-biased and unbiased genes had significantly higher ω values than male-biased (Wilcoxon rank sum test, P = 0.0101, P = 0.0146, respectively). However, there was no significant difference in ω values between female-biased and unbiased genes (Wilcoxon rank sum test, P = 0.2887). Since the number and evolutionary rates of male-biased genes in mature flowers are lower than those in flora buds, we decided to focus on the latter in subsequent analyses.

Revised version: In free-ratios model, ω values are free to vary in each branch compared to fixed-ratio branch model and two-ratio branch model. The ‘free-ratio’ branch model yielded interesting results. In floral buds, the median ω values for female-biased, male-biased, and unbiased genes were 0.222, 0.265 and 0.226, respectively (Fig. 4C and Table S5). Male-biased genes had a significantly higher median relative to female-biased genes (19.37% higher, Wilcoxon rank sum test, P = 0.0009) and unbiased genes (17.26% higher, Wilcoxon rank sum test, P = 0.0004) in floral buds. However, there was no significant difference in ω values between female-biased and unbiased genes (Wilcoxon rank sum test, P = 0.9862). In flowers at anthesis, the median ω values for female-biased, male-biased, and unbiased genes were 0.300, 0.148 and 0.227, respectively (Fig. 4D and Table S5). Female-biased and unbiased genes had significantly higher ω values than male-biased (Wilcoxon rank sum test, P= 0.0101, P = 0.0146, respectively). However, there was no significant difference in ω values between female-biased and unbiased genes (Wilcoxon rank sum test, P = 0.2887). Since the number of male-biased genes and evolutionary rates of male-biased genes in flowers at anthesis are lower than those in floral buds, we decided to focus on the latter in subsequent analyses.

Earlier version: Evidences and contributions of positive selection and relaxed selection for male-biased genes in floral buds

Revised version: Evidence of positive selection and relaxed selection for male-biased genes in floral buds

Earlier version: After comparing the A model and null model, we discovered that 39 out of 343 OGs (11.34%) exhibited strong evidence of having certain sites that evolved under positive selection based on foreground 2b ω value, likelihood ratio tests (LRTs, P < 0.05) and Bayes empirical Bayes (BEB) value (Fig. 5 and Table S6). As an additional approach, we utilized the aBSREL and BUSTED methods that were implemented in HyPhy v.2.5 software, and detected significant evidence of positive selection. According to our findings, 84 out of 343 OGs (24.49%) were identified to be under episodic positive selection in male-biased genes of floral buds with a site proportion of 0.17%–26.44% based on aBSREL (Table S7). In addition, 69 out of 343 OGs (20.01%) exhibited significant signs of positive selection with the site proportion of 0.28%–32.65% in male-biased genes of floral buds according to BUSTED (Table S8). Among these, a total of 32 OGs (9.30%) were identified through our tests using CodeML, aBSREL and BUSTED (Fig. 5).

Revised version: After comparing the A model and null model (see Methods Section), we discovered that 39 out of 343 OGs (11.34%) in male-biased genes of floral buds exhibited strong evidence of having certain sites that evolved under positive selection based on foreground 2b ω value, likelihood ratio tests (LRTs, P < 0.05) and Bayes empirical Bayes (BEB) value (Fig. 5 and Table S6). As an additional approach, we utilized the aBSREL and BUSTED methods that were implemented in HyPhy v.2.5 software, and detected significant evidence of positive selection. According to our findings, 84 out of 343 OGs (24.49%) were identified to be under episodic positive selection in male-biased genes of floral buds with a site proportion of 0.17%–26.44% based on aBSREL (Table S7). In addition, 69 out of 343 OGs (20.01%) exhibited significant signs of positive selection with the site proportion of 0.28%–32.65% in male-biased genes of floral buds according to BUSTED (Table S8). Among these, a total of 32 OGs (9.30%) were identified through our tests using CodeML, aBSREL and BUSTED (Fig. 5).

Venn diagrams of male-biased genes detected to be under positive selection using aBSREL, BUSTED, CodeML and RELAX in floral buds.

Earlier version: Relaxed selection may occur when the efficiency of natural selection is dramatically reduced or eliminated at both gene-wide and genome-wide levels. This has been proposed as an explaination for the rapid evolution of sex-biased genes (Lahti et al., 2009; Mank, 2017). Using the RELAX model, we conducted an analysis and found that 18 out of 343 OGs (5.23%) showed significant evidence of relaxed selection (K = 0.0184–0.6497) (Table S9). Interestingly, we also observed that 61 out of 343 OGs (17.73%) exhibited significant evidence of intensified positive selection (K = 2.3363–50, ω2 ≥ 1) (Fig. 5 and Table S10), which is consistent with the results obtained from CodeML, aBSREL and BUSTED.

Revised version: Relaxed selection may occur when the efficiency of natural selection (e.g., the reduction of the strength of purifying selection) is dramatically reduced or eliminated at both gene-wide and genome-wide levels, leading to accumulations of deleterious mutations relative to positive selection (Lahti et al., 2009; Dapper and Wade, 2020). This has been proposed as an explanation for the rapid evolution of sex-biased genes (Lahti et al., 2009; Mank, 2017). Using the RELAX model, we conducted an analysis and found that 18 out of 343 OGs (5.23%) showed significant evidence of relaxed selection (K = 0.0184–0.6497) (Tables S9 and S13). Interestingly, we also observed that 61 out of 343 OGs (17.73%) exhibited significant evidence of intensified positive selection (K = 2.3363–50, ω2 ≥ 1) (Fig. 5 and Table S10), which is consistent with the results obtained from CodeML, aBSREL and BUSTED.

Earlier version: According to previous studies (Ellegren and Parsch, 2007; Catalan et al., 2018), genes that exhibit sex-biased expression with rapid evolutionary rates tend to display a lower codon bias usage compared to unbiased genes. In our results, we found that male-biased genes in floral buds had a significantly lower median effective number of codons (ENCs) than both female-biased and unbiased genes (Wilcoxon rank sum test, female-biased vs male-biased genes, P = 0.0001 and male-biased vs unbiased genes, P = 0.0123). This suggested that male-biased genes in floral buds exhibit a stronger codon bias usage than both female-biased and unbiased genes (Fig. S3). In summary, our analyses indicated that rapid evolutionary rates of male-biased genes in floral buds were not associated with a reduction in codon usage bias.

Revised version: According to previous studies (Ellegren and Parsch, 2007; Catalan et al., 2018), genes that exhibit sex-biased expression with rapid evolutionary rates tend to display a lower codon bias compared to unbiased genes. In our results, we found that male-biased genes in floral buds had a significantly lower median effective number of codons (ENCs) than both female-biased and unbiased genes (Wilcoxon rank sum test, female-biased vs male-biased genes, P = 0.0001 and male-biased vs unbiased genes, P = 0.0123). This suggested that male-biased genes in floral buds exhibit stronger codon bias than both female-biased and unbiased genes (Fig. S3). In short, our analyses indicated that rapid evolutionary rates of male-biased genes in floral buds were not associated with a reduction in codon usage bias.

Earlier version: To confirm the contributions of positive selection and relaxed selection to rapid rates of male-biased genes in floral buds, we generated three datasets of OGs by excluding different sets of genes. Specifically, we excluded 18 relaxed selective male-biased genes (5.23%), 98 positively selected male-biased genes (28.57%), and 112 male-biased genes (32.65%) under positive and relaxed selection from 343 OGs (Fig. S4). We observed that after excluding relaxed selective male-biased genes, the median (0.264) decreased by 0.34% compared to the median (0.265) of all OGs (Fig. S4A-B). However, after excluding positively selected male-biased genes, the median (0.236) reduced by 11% (Fig. S4A, C) in the results of ‘free-ratio’ branch model. This pattern was consistene in the results of ‘two-ratio’ branch model as well (Fig. S4E-G). Additionally, we analyzed female-biased and unbiased genes that underwent positive and relaxed selection in floral buds (Tables S6-S10). We identified 216 (18.86%) positively selected (Fig. S5A), and 69 (6.03%) relaxed selective female-biased genes from 1,145 OGs (Fig. S5B), respectively. Similarly, we found 436 (18.33%) positively selected (Fig. S6A), and 43 (1.81%) unbiased genes under relaxed selection from 2,378 OGs (Fig. S6B), respectively. Notably, male-biased genes have a higher proportion (10%) of positively selected genes compared to female-biased and unbiased genes. However, relaxed selective male-biased genes have a higher proportion (3.24%) than unbiased genes, but about 0.8% lower than that of female-biased genes. In summary, our analyses suggested that positive selection and relaxed selection likely drove the rapid evolutionary rates of male-biased genes compared to female-biased and unbiased genes in floral buds, but positive selection had an even greater impact.

Revised version: To confirm the contributions of positive selection and relaxed selection to rapid rates of male-biased genes in floral buds, we generated three datasets of OGs by excluding different sets of genes. Specifically, we excluded 18 relaxed selective male-biased genes (5.23%), 98 positively selected male-biased genes (28.57%), and 112 male-biased genes (32.65%) under positive and relaxed selection from 343 Ogs (Fig. S4). We observed that after excluding male-biased genes under relaxed purifying selection, the median (0.264) decreased by 0.34% compared to the median (0.265) of all OGs (Fig. S4A-B). However, after excluding positively selected male-biased genes, the median (0.236) was reduced by 11% (Fig. S4A, C) in the results of ‘free-ratio’ branch model. This pattern was consistent with the results of ‘two-ratio’ branch model as well (Fig. S4E-G). Additionally, we analyzed female-biased and unbiased genes that underwent positive and relaxed selection in floral buds (Tables S6-S10). We identified 216 (18.86%) positively selected (Fig. S5A), and 69 (6.03%) relaxed selective female-biased genes from 1,145 OGs (Fig. S5B), respectively. Similarly, we found 436 (18.33%) positively selected (Fig. S6A), and 43 (1.81%) unbiased genes under relaxed selection from 2,378 OGs (Fig. S6B), respectively. Notably, male-biased genes have a higher proportion (10%) of positively selected genes compared to female-biased and unbiased genes. However, relaxed selective male-biased genes have a higher proportion (3.24%) than unbiased genes, but about 0.8% lower than that of female-biased genes. In summary, our analyses suggested that positive selection and relaxed selection likely drove the rapid evolutionary rates of male-biased genes compared to female-biased and unbiased genes in floral buds, but positive selection had an even greater impact.

Functional analysis of sex-biased genes in floral buds

We conducted KEGG pathway enrichment analysis on sex-biased genes in floral buds. Our results showed that 699 genes were female-biased and 358 genes were male-biased, with significant enrichment (P < 0.05) in 26 and 24 KEGG pathways, respectively (Table S11). In the floral bud stage, we observed that female-biased genes were mainly enriched in metabolic and signaling pathways, such as ribosome, Fatty acid elongation, photosynthesis and plant hormone signal transduction (Fig. S7A and Table S11). On the other hand, male-biased genes were significantly enriched in metabolic and signaling pathways, including inositol phosphate metabolism, starch and sucrose metabolism, regulation of autophagy, plant hormone signal transduction, and Toll-like receptor signaling pathway (Fig. S7B and Table S11).

Earlier version: We have also found that certain male-biased genes, which are evolving under positive selection and relaxed selection (Table S12 and S13), were related to abiotic and biotic stress. For instance, mitogen-activated protein kinase kinase kinase 18 (MAPKKK18) (Zhang and Zhang 2022,), zinc finger CCCH domain-containing protein 20 (C3H20/TZF2) (Bogamuwa and Jang, 2014), and heat stress transcription factor B-3 (HSFB3) (Scharf et al., 2012) have been linked to stress. Interestingly, ten male-biased genes with rapid evolutionary rates were associated with anther and pollen development. These genes include LRR receptor-like serine/threonine protein kinase (LRR-RLK) (Cui et al., 2022), pollen receptor-like kinase 3 (PRK3) (Muschietti and Wengier, 2018), autophagy-related protein 18f (ATG18f) (Zhou et al., 2015; Li et al., 2020a), and plant homeodomain (PHD) finger protein 3 (MALE STERILITY 3) (Hou et al., 2022) in floral buds of male plants.

Revised version: We have also found that certain male-biased genes, which are evolving under positive selection and relaxed selection (Table S12 and S13), were related to abiotic stress and immune responses. For instance, mitogen-activated protein kinase kinase kinase 18 (MAPKKK18) (Zhang and Zhang, 2022), zinc finger CCCH domain-containing protein 20 (C3H20/TZF2) (Bogamuwa and Jang, 2014), and heat stress transcription factor B-3 (HSFB3) (Scharf et al., 2012) have been linked to stress. Additionally, ten male-biased genes with rapid evolutionary rates were associated with anther and pollen development. These genes include LRR receptor-like serine/threonine protein kinase (LRR-RLK) (Cui et al., 2022), pollen receptor-like kinase 3 (PRK3) (Muschietti and Wengier, 2018), autophagy-related protein 18f (ATG18f) (Zhou et al., 2015; Li et al., 2020), and plant homeodomain (PHD) finger protein 3 (MALE STERILITY 3) (Hou et al., 2022) in floral buds of male plants.

Discussion

Earlier version: The Cucurbitaceae family is an excellent model for studying the evolution of sexual systems of angiosperms, including sex-determination mechanism and sexual dimorphism. In this study, we aimed to investigate whether sex-biased genes exhibited the evidence of rapid evolutionary rates of protein sequences and identify the evolutionary forces responsible for the observed patterns in the dioecious T. pilosa. We compared the expression profiles of sex-biased gene between sexes and two tissue types and examined the signatures of rapid sequence evolution for sex-biased genes, as well as the contributions of potential evolutionary forces.

Revised version: The Cucurbitaceae family, where half of the species are monoecious and half are dioecious, is an excellent model for studying the evolution of sexual systems of angiosperms, including sex-determination mechanism and sexual dimorphism (Schaefer and Renner, 2010; Boualem et al., 2015; Ma and Pannell, 2016). In this study, we aimed to investigate whether sex-biased genes exhibited evidence of rapid evolutionary rates of protein sequences and identify the evolutionary forces responsible for the observed patterns in the dioecious Trichosanthes pilosa. We compared the expression profiles of sex-biased genes between sexes and two tissue types and examined the signatures of rapid sequence evolution for sex-biased genes, as well as the contributions of potential evolutionary forces.

Sex-biased expression in floral buds

Earlier version: Numerous studies have shown that in dioecious plants, male-biased genes tend to outnumber female-biased genes. For instance, insect-pollinated dioecious plants such as Asparagus officinalis (Harkess et al., 2015) and Silene latifolia (Zemp et al., 2016), exhibit a higher proportion of male-biased genes, whereas wind-pollinated dioecious plants like Populus balsamifera (Sanderson et al., 2019) has twice as many female-biased genes than male-biased genes. These differences in these studies could be partly attributed to the impact of sexual selection on secondary sexual traits in insect-pollinated dioecious plants, as opposed to wind-pollinated ones (Delph and Herlihy, 2012; Muyle, 2019; Sanderson et al., 2019). In contrast to the above studies, our findings revealed that the number of female-biased genes in floral buds of the night-flowering, insect-pollinated dioecious plant Trichosanthes pilosa exceeded that of male-biased genes by 882. This excess of female-biased expression could be due to lower energy consumption needs and reduced chemical defense capability against herbivores in short-lived male flowers (Sanderson et al., 2019). Indeed, functional enrichment analysis in chemical defense pathways such as terpenoid backbone and diterpenoid biosynthesis, indicated that female floral buds had more expressed genes and were better equipped to defend against herbivores compared to male floral buds (Fig. S7A and Table S11). Additionally, our enrichment analysis showed that the photosynthesis, porphyrin and chlorophyll metabolism pathways were more active in female floral buds compared to male floral buds (Fig. S7A and Table S11), enabling them to acquire more resources such as carbon for fruit and seed production (Delph,1999).

Revised version: Numerous studies have shown that in dioecious plants, male-biased genes tend to outnumber female-biased genes, consistent with the patterns in most animals (Zhang et al., 2007; Djordjevic et al., 2022). For instance, insect-pollinated dioecious plants such as Asparagus officinalis (Harkess et al., 2015) and Silene latifolia (Zemp et al., 2016), exhibit a higher proportion of male-biased genes, whereas wind-pollinated dioecious plants like Populus balsamifera (Sanderson et al., 2019) has twice as many female-biased genes than male-biased genes. The differences in these studies could be partly attributed to the impact of sexual selection on secondary sexual traits in insect-pollinated dioecious plants, as opposed to wind-pollinated ones (Delph and Herlihy, 2012; Muyle, 2019; Sanderson et al., 2019). In contrast to the above studies, our findings revealed that the number of female-biased genes in floral buds of the night-flowering, insect-pollinated dioecious plant Trichosanthes pilosa exceeded that of male-biased genes by 882 (∼21%). This excess of female-biased expression could be due to lower energy consumption needs and reduced chemical defense capability against insect herbivores in short-lived male flowers (Sanderson et al., 2019). Indeed, functional enrichment analysis in chemical pathways such as terpenoid backbone and diterpenoid biosynthesis indicated that relative to male floral buds, female floral buds had more expressed genes that were equipped to defend against herbivorous insects and pathogens, except for growth and development (Vaughan et al., 2013; Ren et al., 2022) (Fig. S7A and Table S11). Additionally, our enrichment analysis showed that the photosynthesis, porphyrin and chlorophyll metabolism pathways were more active in female floral buds compared to male floral buds (Fig. S7A and Table S11), enabling them to acquire more resources such as carbon for fruit and seed production (Delph,1999).

Earlier version: We identified functional enrichments in Toll-like receptor signaling, NF-kappa B signaling, and Inositol phosphate metabolism pathways in male floral buds (Fig. S7B, Tables S11 and S12). These findings suggest that male floral buds are better adaptated to biotic and abiotic environments. Moreover, the enrichment in regulation of autophagy pathways could be associated with the senescence for male floral buds (Table S14) (Liu and Bassham, 2012; Li et al., 2020). We also observed that homologous genes of two male-biased genes (Table S14) that control the raceme inflorescence development (Teo et al., 2014) were highly expressed in floral buds. This suggests that sex-biased genes, rather than sex-specific genes, are responsible for dimorphic traits of inflorescence. In short, these results indicate that changes in sex-biased genes expression play different roles on sexual dimorphism in physiological and morphological traits (Dawson and Geber, 1999).

Revised version: We identified functional enrichments in Toll-like receptor signaling, NF-kappa B signaling, and inositol phosphate metabolism pathways in male floral buds (Fig. S7B, Table S11). We also found that male-biased genes with high evolutionary rates in male floral buds were associated with functions to abiotic stresses and immune responses (Tables S12 and S13), which suggest that male floral buds through rapidly evolving genes are adapted to mountain climate and the environment in Southwest China relative to female floral buds through high gene expression. In addition, the enrichment in regulation of autophagy pathways could be associated with gamete development and the senescence of male floral buds (Table S14) (Liu and Bassham, 2012; Li et al., 2020; Zhou et al., 2021). In fact, it was observed that male flowers senesced faster (Wu et al., 2011). We also found that homologous genes of two male-biased genes in floral buds (Table S14) that control the raceme inflorescence development (Teo et al., 2014) were highly expressed compared to female floral buds. Taken together, these results indicate that expression changes in sex-biased genes, rather than sex-specific genes play different roles in sexual dimorphic traits in physiology and morphology (Dawson and Geber, 1999).

Rapid Evolution of male-biased genes in floral buds

Earlier version: It has been observed that, in most animals, sex-biased genes, particularly those biased towards males, often exhibit more rapid evolutionary rates than unbiased genes (Parsch and Ellegren, 2013; Grath and Parsch, 2016; Mank, 2017; Toubiana et al., 2021). However, in dioecious angiosperms, no evidence of rapid evolution in sex-biased genes relative to unbiased genes has been found (Zemp et al., 2016; Darolti et al., 2018; Cossard et al., 2019; Sanderson et al., 2019; Scharmann et al., 2021). In contrast, our findings indicated that male-biased genes exhibited rapid evolutionary rates than both female-biased and unbiased genes in floral buds of dioecious T. pilosa. We proposed that positive selection and relaxed purifying selection may be responsible for the rapid sequence evolution of male-biased genes.

Revised version: It has been observed that, in most animals, sex-biased genes, particularly those biased towards males, often exhibit more rapid evolutionary rates than unbiased genes (Parsch and Ellegren, 2013; Grath and Parsch, 2016; Mank, 2017; Toubiana et al., 2021). However, in dioecious angiosperms, no evidence of rapid evolution in sex-biased genes relative to unbiased genes has been found (Zemp et al., 2016; Darolti et al., 2018; Cossard et al., 2019; Sanderson et al., 2019; Scharmann et al., 2021). In contrast, our findings indicated that male-biased genes experience higher evolutionary rates than both female-biased and unbiased genes in floral buds of dioecious T. pilosa. We proposed that positive selection and relaxed purifying selection may be responsible for the rapid sequence evolution of male-biased genes.

Earlier version: After analyzing the data, we found that around 28.57% (98 genes) of male-biased genes have undergone positive selection. Additionally, we observed that the proportion of male-biased genes under positive selection was about 10% higher than that of female-biased and unbiased genes. Furthermore, we discovered that some male-biased genes under positive selection were linked to abiotic and biotic stress responses (Table S12). Our findings are consistent with studies on Drosophila and Ectocarpus (Zhang and Parsch, 2005; Lipinska et al., 2015), suggesting that adaptive evolution is one of the primary evolutionary forces for rapid evolutionary rates. Notably, we identified several male-biased genes under positive selection that are functionally related to early flowering (phyB) (Stephenson and Bertin, 1983; Forrest, 2014; Hajdu et al., 2015) and pollen development (Skogsmyr and Lankinen, 2002; Williams and Reese, 2019) (Tables S12-S14). These findings indicate that a small fraction of male-biased genes may experience adaptive evolution due to sexual selection, driven by male-male competition.

Revised version: After analyzing the data, we found that around 28.57% (98 genes) of male-biased genes have undergone positive selection. Additionally, we observed that the proportion of male-biased genes under positive selection was about 10% higher than that of female-biased and unbiased genes. Furthermore, we discovered that some male-biased genes under positive selection were linked to abiotic stress and immune responses (Table S12). Our findings are consistent with studies on Drosophila and Ectocarpus (Zhang and Parsch, 2005; Lipinska et al., 2015), suggesting that adaptive evolution is one of the important driving forces for rapid evolutionary rates. Notably, we identified several male-biased genes under positive selection that are functionally related to early flowering (phyB) (Stephenson and Bertin, 1983; Forrest, 2014; Hajdu et al., 2015) and pollen development (Skogsmyr and Lankinen, 2002; Williams and Reese, 2019) (Tables S12- S14). These findings indicate that a small fraction of male-biased genes may experience adaptive evolution due to sexual selection, driven by male-male competition.

Earlier version: Alternatively, relaxed constraints could contribute to the rapid evolutionary rates of sex-biased genes through three key characteristics (Dapper and Wade, 2020; Tosto et al., 2023). First, sex-biased genes are often expressed solely in reproductive tissues of one sex (e.g., sex-specific genes), particularly in the haploid phase (Immler, 2019; Beaudry et al., 2020). Sex-specific selection (e.g., relaxed purifying selection) acting on sex-specific genes could decrease the elimination of deleterious mutations (Mank, 2017), such as pollen-specific (Harrison et al., 2019) or testes-specific genes (Gershoni and Pietrokovski, 2014). However, our results did not reveal male-specific genes that underwent relaxed purifying selection. Second, sex-biased genes are often expressed in few tissues (tissue-specific genes) (Meisel, 2011; Tosto et al., 2023). These genes could evolve under positive selection or relaxed purifying selection due to low pleiotropy (Congrains et al., 2018; Whittle et al., 2021; Tosto et al., 2023). In our results, 343 male-biased genes with faster evolutionary rates relative to female-biased and unbiased genes overlapped with tissue-biased genes in floral buds (27 out of 343, 7.87%) (Fig. S8A). Furthermore, 27 out of 343 male-biased genes (that is, tissue-biased genes) in floral buds overlapped with nine out of 98 (9.18%) male-biased genes under positive selection (Fig. S8B), and one out of 18 (5.56%) male-biased genes under relaxed purifying selection (Fig. S8C). Consistent with male-biased genes in Fucus (Hatchett et al., 2023), our analyses suggested that elevated evolutionary rates may partly be linked to low pleiotropy. Finally, gene duplication has long been thought to promote functional divergences and phenotypic novelties by relaxing the constraints of purifying selection on the duplicated gene copy early in its history (Lynch and Conery, 2000; Lynch and Katju, 2004; Lahti et al., 2009). For instance, the progesterone receptor gene family in the human lineage (Marinić and Lynch, 2020) and the CYP98A9 clade in Brassicales (Liu et al., 2016) have demonstrated rapid evolution and divergent function due to relaxed purifying selection. In our results, we identified 18 out of 343 (5.25%) male-biased genes that underwent relaxed purifying selection using RELAX model (Table S13). Interestingly, the vast majority of genes under relaxed selection were members of different gene families generated by gene duplication (including whole-genome duplication), such as LOB domain-containing protein 18 (LBD18) (Zhang et al., 2020), WRKY transcription factor 72 (WRKY72) (Chen et al., 2017), and pollen receptor-like kinase 3 (PRK3) (Muschietti and Wengier, 2018).

Revised version: Alternatively, relaxed constraints could contribute to the rapid evolutionary rates of sex-biased genes through three key characteristics (Dapper and Wade, 2020; Tosto et al., 2023). First, sex-biased genes are often expressed solely in reproductive tissues of one sex (e.g., sex-specific genes), particularly in the haploid phase (Immler, 2019; Beaudry et al., 2020). Sex-specific selection (e.g., relaxed purifying selection) acting on sex-specific genes could decrease the elimination of deleterious mutations (Mank, 2017), such as pollen-specific (Harrison et al., 2019) or testes-specific genes (Gershoni and Pietrokovski, 2014). However, our results observed male-biased genes but not male-specific genes that underwent relaxed purifying selection. Second, sex-biased genes are often expressed in few tissues (tissue-biased genes) (Meisel, 2011; Tosto et al., 2023), resulting in these genes rapidly evolving under positive selection or under relaxed purifying selection due to low pleiotropic constraints (Congrains et al., 2018; Whittle et al., 2021; Tosto et al., 2023). In our results, 343 male-biased genes with faster evolutionary rates relative to female-biased and unbiased genes overlapped with tissue-biased genes in floral buds (27 out of 343, 7.87%) (Fig. S8A). Furthermore, 27 out of 343 male-biased genes (that is, tissue-biased genes) in floral buds overlapped with nine out of 98 (9.18%) male-biased genes under positive selection (Fig. S8B), and one out of 18 (5.56%) male-biased genes under relaxed purifying selection (Fig. S8C). So, we obtained ten rapidly evolving tissue-biased genes which were also male-biased in male flower buds, suggesting that elevated evolutionary rates may partly be linked to low pleiotropy, consistent with male-biased genes in Fucus (Hatchett et al., 2023). Finally, gene duplication has long been thought to promote functional divergences and phenotypic novelties by relaxing the constraints of purifying selection on the duplicated gene copy early in its history (Lynch and Conery, 2000; Lynch and Katju, 2004; Lahti et al., 2009). For instance, the progesterone receptor gene family in the human lineage (Marinić and Lynch, 2020) and the CYP98A9 clade in Brassicales (Liu et al., 2016) have demonstrated rapid evolution and divergent function due to relaxed purifying selection. In our results, we identified 18 out of 343 (5.25%) male-biased genes that underwent relaxed purifying selection using RELAX model (Table S13). Interestingly, the vast majority of genes under relaxed selection were members of different gene families generated by gene duplication (including whole-genome duplication), such as LOB domain-containing protein 18 (LBD18) (Zhang et al., 2020), WRKY transcription factor 72 (WRKY72) (Chen et al., 2017), and pollen receptor-like kinase 3 (PRK3) (Muschietti and Wengier, 2018).

Earlier version: Additionally, although reducing codon usage bias could theoretically accelerate evolutionary rates of sex-biased genes by decreasing synonymous substitution rates, our results did not support this idea. In fact, we found that male-biased genes had significantly higher median dS values than female-biased and unbiased genes, both in the ‘free-ratio’ analysis (Fig. S9A, female-biased vs male-biased genes, P = 6.444e-12 and male-biased vs unbiased genes, P = 4.564e-13) and ‘two-ratio’ branch model (Fig. S9B, female-biased vs. male-biased genes, P = 2.2e-16 and male-biased vs. unbiased genes, P = 9.421e-08, respectively), indicating higher synonymous substitution rates.

Revised version: Additionally, reducing codon usage bias could theoretically accelerate evolutionary rates of sex-biased genes by decreasing synonymous substitution rates. However, our results did not support this idea due to stronger codon usage bias in male-biased genes. Codon usage bias is influenced by many factors, such as levels of gene expression. Highly expressed genes have a stronger codon usage bias and could be encoded by optimal codons for more efficient translation (Frumkin et al., 2018; Parvathy et al., 2022), consistent with high levels of gene expression in males (that is, male-biased genes) in floral buds. Additionally, stronger codon usage bias may be related to higher synonymous substitution rates (Parvathy et al., 2022). Indeed, male-biased genes had significantly higher median dS values than female-biased and unbiased genes, both in the ‘free-ratio’ analysis (Fig. S9A, female-biased vs male-biased genes, P = 6.444e-12 and male-biased vs unbiased genes, P = 4.564e-13) and ‘two-ratio’ branch model (Fig. S9B, female-biased versus male-biased genes, P = 2.2e-16 and male-biased versus unbiased genes, P = 9.421e-08, respectively). Hence, rapid evolution of male-biased genes in floral buds was not associated with codon usage bias in our results.

Earlier version: Several species have been observed to exhibit rapid evolutionary rates of sequences on sex chromosome compared to autosomes, which has been related to the evolutionary theories of fast-X or fast-Z (Meisel and Connallon, 2013; Wright et al., 2015; Charlesworth et al., 2018). Furthermore, the quantification of gene expression by bulk RNA-seq technology, relative to single-cell transcriptome analysis, has been shown to potentially obfuscating true signals in the evolution of sex-biased gene expression in complex aggregations of diverse cell types (Darolti and Mank, 2023; Tosto et al., 2023). However, investigation of these interesting issues related to sex-biased gene evolution in T. pilosa can only be conducted when whole genome sequences become available in the near future.

Revised version: Several species have been observed to exhibit rapid evolutionary rates of sequences on sex chromosomes compared to autosomes, which has been related to the evolutionary theories of fast-X or fast-Z (Hough et al., 2014; Wright et al., 2015; Charlesworth et al., 2018a). Furthermore, the quantification of gene expression by bulk RNA-seq technology, relative to single-cell transcriptome analysis, has been shown to potentially obfuscate true signals in the evolution of sex-biased gene expression in complex aggregations of diverse cell types (Darolti and Mank, 2023; Tosto et al., 2023). However, investigation of these interesting issues related to sex-biased gene evolution in T. pilosa can only be conducted when whole genome sequences become available in the near future.

Methods

Plant materials and RNA isolation

Earlier version: Floral buds (≤ 3 mm) and mature flowers were sampled from three female and three male plants (Fig. 1) in the field of Qinglong Gorge of Anning, Yunnan Province. Floral buds from female and male plants were named F1 and M1, respectively. Similarly, mature flowers from female and male plants, there were named F2 and M2, respectively (Table S1). To exclude possible bacterial contamination, all tissues were sterilized with 75% alcohol and immediately rinsed with purified water. All samples were then snap-frozen in liquid nitrogen, and stored at −80°C. Total RNA was extracted from each sample using TRIzol reagent (Life Technologies, CA, USA) according to the manufacturer’s instructions. The quantification and qualification of RNA were assessed by the RNA Nano 6000 Assay Kit of the Bioanalyzer 2100 system (Agilent Technologies, CA, USA).

Revised version: Floral buds (≤ 3 mm) and flowers at anthesis were sampled from three female and three male plants (Fig. 1) from the mountainous regions of Anning (Qinglong Gorge), Yunnan Province in Southwest China. Floral buds from female and male plants were named F1 and M1, respectively. Similarly, flowers at anthesis from female and male plants, there were named F2 and M2, respectively (Table S1). To exclude possible bacterial contamination, all tissues were sterilized with 75% alcohol and immediately rinsed with purified water. All samples were then snap-frozen in liquid nitrogen, and stored at -80°C. Total RNA was extracted from each sample using TRIzol reagent (Life Technologies, CA, USA) according to the manufacturer’s instructions. The quantification and qualification of RNA were assessed by the RNA Nano 6000 Assay Kit of the Bioanalyzer 2100 system (Agilent Technologies, CA, USA).

Illumina sequencing, de novo assembly and annotation

To construct the library, approximately 2 μg of total RNA was used with the Illumina NEBNext UltraTM RNA Library Prep Kit. RNA sequencing was performed on the Illumina NovaSeq 6000, generating 150 bp paired-end reads. The resulting clean reads were obtained by removing adapters, reads containing N bases and low-quality reads using Trimmomatic v.0.39 (Bolger et al., 2014). These reads were deposited in the NCBI database (PRJNA899312).

De novo assembly for clean reads from all samples was performed using Trinity v.2.10.0 (Haas et al., 2013) with min_kmer_cov: 3 and all other default parameters. To eliminate contamination, all transcripts of de novo assembly were compared to bacterial genomes downloaded from NCBI databases using BLASTN with e-value of 1.0e-05 in blast+ 2.12.0 software. We used Corset v.4.6 (Davidson and Oshlack, 2014) to obtain high quality, non-redundant consensus transcripts (unigenes). TransDecoder v.5.5.0 was run with -m 100 parameters, namely at least 100 amino acids, to predict the coding DNA and protein sequences (Haas et al., 2013).

Earlier version: To evaluate the accuracy and completeness of reference transcriptomes, we performed gene function annotations based on the following databases, using BLAST with a cutoff e-value of 1.0e-05: NR, NT and Swissport (Shiryev et al., 2007). We mapped the unigenes to Pfam database using InterProScan v.5.41 (Jones et al., 2014), to the GO database using Blast2GO (Conesa et al., 2005), and to the KEGG database using KEGG automatic annotation server (Moriya et al., 2007). Additionally, we estimated the completeness of reference transcriptomes using BUSCO v.5.2.2 based on embryophyta_odb9 database (Seppey et al., 2019).

Revised version: To evaluate the accuracy and completeness of reference transcriptomes, we performed gene function annotations based on the following databases, using BLAST with a cutoff e-value of 1.0e-05: NR, NT and Swissport (Shiryev et al., 2007). We mapped the unigenes to Pfam database using InterProScan v.5.41 (Jones et al., 2014), to the GO database using Blast2GO (Conesa et al., 2005), and to the KEGG database using KEGG automatic annotation server (Moriya et al., 2007). Additionally, we estimated the completeness of reference transcriptomes using BUSCO v.5.4.5 based on embryophyta_odb10 database (Seppey et al., 2019).

Detection of sex-biased genes

Earlier version: Clean reads were mapped onto all unigenes using Bowtie2 (Langmead and Salzberg, 2012). Read counts were normalized to FPKM (Fragments Per Kilobase Million) value for each unigene using RSEM (Li and Dewey, 2011) in different male and female samples. Genes with zero read counts (i.e., no expression) in both two sexes and tissues were excluded. Differential expression analysis between sexes and tissue types was performed using DESeq2 R package. Unigenes with an FDR-adjusted P < 0.05 and an absolute value of log2 ratio ≥ 1 identified by DESeq2 were considered as sex-biased genes. To perform KEGG functional enrichment, we use all KEGG annotation terms for all genes as the background and performed the analyses using KOBAS v.2.0.12 (Mao et al., 2005).

Revised version: Clean reads were mapped onto all unigenes using Bowtie2 (Langmead and Salzberg, 2012). Read counts were normalized to FPKM (Fragments Per Kilobase Million) value for each unigene using RSEM (Li and Dewey, 2011) in different male and female samples. Genes with zero read counts (i.e., no expression) in both two sexes and tissues were excluded. Differential expression analysis between sexes and tissue types was performed using DESeq2 R package (Love et al., 2014). Unigenes with an FDR-adjusted P < 0.05 and an absolute value of log2 ratio ≥ 1 identified by DESeq2 were considered as sex-biased genes. To perform KEGG functional enrichment, we used all KEGG annotation terms for all genes as the background and performed the analyses using KOBAS v.2.0.12 (Mao et al., 2005).

Evolutionary rate analyses

To quantify the evolutionary rates of sex-biased genes, we download published genome datasets for monoecious Trichosanthes anguina (Ma et al., 2020) and monoecious Luffa cylindrica which has a closer phylogenetic relationship with Trichosanthes (de Boer et al., 2012; Wu et al., 2020) from CuGenDB database (Zheng et al., 2019). Additionally, we also download published RNA sequencing reads of floral buds and flowers from CNCB (Accession CRA002313) and NCBI databases (Accession SRR5259239) for dioecious plant Trichosanthes kirilowii (Hu et al., 2020), and de novo assembled by previously described methods.

We identified one-to-one OGs using OrthoFinder v.2.3.3 with default parameters from T. anguina, T. pilosa, T. kirilowii, and Luffa cylindrica (Emms and Kelly, 2019). Then, we employed TranslatorX with -c 1 -p M -g -b5 n parameters (i.e., the multiple alignment and the trimming using Muscle and GBlocks, respectively), translated nucleotide sequences and back-translated amino acid alignments into nucleotide alignments to ensure codon-to-codon alignment (Abascal et al., 2010). The remaining gapless alignments (≥100 bp in length) were retained.

Earlier version: To investigate the evolutionary rates of code sequences, we estimated nonsynonymous substitution (dN), synonymous substitution (dS) rates, as well as protein substitution rates (dN/dS, ω), using two branch models from CodeML package in PAML v.4.9h with the F3x4 codon frequencies (CodonFreq = 2) (Yang, 2007). According to the phylogenetic relationships of Trichosanthes (de Boer et al., 2012; Guo et al., 2020), we set up tree structure ((T. anguina, T. pilosa), T. kirilowii, L. cylindrica) in the control file of CodeML. First, we employed a ‘two-ratio’ branch model (model = 2, Nssites = 0) that assumes the foreground (two dioecious species) has a different ω value from the background (two monoecious species) to estimate and compare the divergences of the foreground. Second, to reduce the potential bias of ω value due to the conflation of two dioecious species, we also implemented a ‘free-ratio’ branch model (model = 1, Nssites = 0), which assumes an independent ω ratio for each branch. Finally, we excluded all OGs with ω > 2, plotted the distribution of ω values and compared the median of ω values in female-biased, male-biased and unbiased orthologous genes of floral buds and mature flowers. All comparisons between sex-biased and unbiased genes were tested using Wilcoxon rank sum test in R software.

Revised version: To investigate the evolutionary rates of coding sequences, we estimated nonsynonymous substitution (dN), synonymous substitution (dS) rates, as well as protein substitution rates (dN/dS, ω), using two branch models from CodeML package in PAML v.4.9h with the F3x4 codon frequencies (CodonFreq = 2) (Yang, 2007). According to the phylogenetic relationships of Trichosanthes (de Boer et al., 2012; Guo et al., 2020), we set up tree structures ((T. anguina, T. pilosa), T. kirilowii, L. cylindrica) in the control file of CodeML. First, we employed a ‘two-ratio’ branch model (model = 2, Nssites = 0) that assumes the foreground (two dioecious species) has a different ω value from the background (two monoecious species) to estimate and compare the divergences of the foreground. Second, to reduce the potential bias of ω value due to the conflation of two dioecious species, we also implemented a ‘free-ratio’ branch model (model = 1, Nssites = 0), which assumes an independent ω ratio for each branch. Finally, to avoid the effects of saturation substitution, we used separately OGs with 0 < ω < 2 and all OGs with ω > 0, plotted the distribution of ω values and compared the median of ω values in female-biased, male-biased and unbiased orthologous genes of floral buds and flowers at anthesis. All comparisons between sex-biased and unbiased genes were tested using Wilcoxon rank sum test in R software.

Estimation of the strength of natural selection

Earlier version: The rapid evolutionary rates of sex-biased genes may be attributed by positive selection, relaxed selection and lower codon usage bias (Catalan et al., 2018; Dapper and Wade, 2020). Therefore, we conducted separate analyses using classical branch-site models, the adaptive branch-site random effects likelihood (aBSREL) model, the branch-site unrestricted statistical test for episodic diversification (BUSTED) model, the RELAX model, and the effective number of codons (ENC) in PAML v.4.9h, HyPhy v.2.5 and CodonW v.1.4.2 to distinguish which evolutionary forces are driving the rapid evolutionary rates of sex-biased genes (Pond et al., 2020).

Revised version: The rapid evolutionary rates of sex-biased genes may be attributed to positive selection, relaxed selection and lower codon usage bias (Catalan et al., 2018; Dapper and Wade, 2020). Therefore, we conducted separate analyses using classical branch-site models, the adaptive branch-site random effects likelihood (aBSREL) model (Smith et al., 2015), the branch-site unrestricted statistical test for episodic diversification (BUSTED) model (Murrell et al., 2015), the RELAX model (Wertheim et al., 2015), and the effective number of codons (ENC) in PAML v.4.9h (Álvarez-Carretero et al., 2023), HyPhy v.2.5 (Pond et al., 2020) and CodonW v.1.4.2 (http://codonw.sourceforge.net) to distinguish which evolutionary forces are driving the rapid evolutionary rates of sex-biased genes.

Earlier version: To determine if amino acid sites in the foreground, including the T. pilosa lineage have undergone positive selection (foreground 2b ω >1) compared with the background for each OGs, we used branch-site model A (model = 2, Nssite = 2, fix_omega = 0, omega = 1.5) and branch-site model null (model = 2, Nssite = 2, fix_omega = 1, omega = 1). We examined the significance of likelihood ratio tests (LRTs, P < 0.05) to identify positively selected sites between model A and model null by comparing LRTs to the Chi-square distribution with two degrees of freedom. We adjusted the LRTs P value for multiple comparisons using Benjamini and Hochberg’s (FDR) algorithm. When the P value was significant, we used Bayes Empirical Bayes (BEB) estimates to identify sites with a high posterior probability (pp ≥ 0.95) of being under positive selection (Yang et al., 2005; Zhang et al., 2005).

Revised version: To determine if amino acid sites in the foreground, including the T. pilosa lineage have undergone positive selection (foreground 2b ω >1) compared with the background for each OGs, we followed the steps of Zhang et al. (2005), and used branch-site model A (model = 2, Nssite = 2, fix_omega = 0, omega = 1.5) and branch-site model null (model = 2, Nssite = 2, fix_omega = 1, omega = 1). We examined the significance of likelihood ratio tests (LRTs, P < 0.05) to identify positively selected sites between model A and model null by comparing LRTs to the Chi-square distribution with two degrees of freedom. We adjusted the LRTs P value for multiple comparisons using Benjamini and Hochberg’s (FDR) algorithm. When the P value was significant, we used Bayes Empirical Bayes (BEB) estimates to identify sites with a high posterior probability (pp ≥ 0.95) of being under positive selection (Yang et al., 2005).

Earlier version: To detect episodic positive selection at a proportion of sites on the foreground branch, we employed the aBSREL method in the HyPhy v.2.5 packages to compare the fully adaptive model (ω > 1) to the null model that allows no positive selection rate classes by LRTs, which is an improved algorithm of branch-site models in PAML (Smith et al., 2015). For relatively small datasets, such as those with fewer than 10 taxa, the aBSREL method may not have enough power to detect positive selection. Therefore, we also ran the BUSTED method to identify gene-wide evidence of episodic positive selection at least one site on at least one branch (Murrell et al., 2015). We set T. pilosa as the foreground and assessed the statistical significance (P < 0.05) using LRTs with the Holm-Bonferroni correction.

Revised version: To detect episodic positive selection at a proportion of sites on the foreground branch, we employed the aBSREL method (Smith et al., 2015) in the HyPhy v.2.5 packages to compare the fully adaptive model (ω > 1) to the null model that allows no positive selection rate classes by LRTs, which is an improved algorithm of branch-site models in PAML. For relatively small datasets, such as those with fewer than 10 taxa, the aBSREL method may not have enough power to detect positive selection. Therefore, we also ran the BUSTED method to identify gene-wide evidence of episodic positive selection at least one site on at least one branch (Murrell et al., 2015). We set T. pilosa as the foreground and assessed the statistical significance (P < 0.05) using LRTs with the Holm-Bonferroni correction.

Earlier version: To test the relaxation of selective strength, we utilized the RELAX model in the HyPhy v.2.5 software (Wertheim et al., 2015; Schrader et al., 2021). The RELAX model estimates three ω parameters (ω0 ≤ ω1 ≤ 1 ≤ ω2), and determines the proportion of sites in the test (foreground) and reference (background) branches using a branch-site model. The first two ω classifications indicate that sites have undergone purifying selection, and the third classification indicates that sites have been under positive selection. Additionally, the model introduces a selection intensity parameter (K value) to compare a null model (K = 1) with an alternative model, thereby assessing the the strength of natural selection. When K > 1, it suggests intensified natural selection, when K < 1, indicates relaxed natural selection in the test branch relative to the reference branch. We quantified the statistical confidence of K value (P < 0.05) using LRTs and the Holm-Bonferroni correction.

Revised version: To test the relaxation of selective strength, we utilized the RELAX model in the HyPhy v.2.5 software (Wertheim et al., 2015; Schrader et al., 2021). The RELAX model estimates three ω parameters (ω0 ≤ ω1 ≤ 1 ≤ ω2), and determines the proportion of sites in the test (foreground) and reference (background) branches using a branch-site model. The first two ω classifications indicate that sites have undergone purifying selection, and the third classification indicates that sites have been under positive selection. Additionally, the model introduces a selection intensity parameter (K value) to compare a null model (K = 1) with an alternative model, thereby assessing the strength of natural selection. When K > 1, it suggests intensified natural selection, when K < 1, indicates relaxed natural selection in the test branch relative to the reference branch. We quantified the statistical confidence of K value (P < 0.05) using LRTs and the Holm-Bonferroni correction.

To investigate codon usage bias, which refers to the differences in the frequency of occurrence of synonymous codons in coding DNA, we employed CodonW v.1.4.2. This program considers the ENC values from 20 to 61 as a measure of the departure of the genetic codes for a given gene (Wright, 1990), with lower ENC values represent stronger codon usage bias (Hambuch and Parsch, 2005). We performed a Wilcoxon rank sum test to determine if there were deviations in ENC values among female-biased, male-biased, and unbiased genes in floral buds.

Supporting information